# ITC - MPSI
---

# Exercices de révision -- vacances de Noël

* Commencez par vérifier que vous avez fait les TP du semestre en entier. Si ce n'est pas le cas, terminez-les.
* Relisez bien la partie en vert en bas des TP ("Les points à retenir") : si un point est flou ou ne vous évoque rien, allez relire la partie cours du TP.
* Je ne corrigerai pas ces exercices : il y a des tests automatiques pour la plupart, et pour les exercices graphiques vous pouvez regarder si ça a l'air bien.
* En général, n'essayez pas de vérifier la nature des arguments, sauf si c'est explicitement demandé dans l'énoncé.

Dans chaque thème, la difficulté est croissante (enfin j'essaie...). À vous de voir où commencer et quoi faire en fonction de vos difficultés. Je reste disponible par mail et sur discord pour toute question.

Les énoncés qui demandent un peu de réflexion (et donc possiblement un papier et un crayon) sont précédés d'étoiles : (*), (\**) ou (\***).

## Fonctions simples et conditionnelles

Écrire une fonction `somme` qui prend en argument deux entiers et renvoie leur somme.

<!-- grade_id: cell-ae86e16a00a83e17, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-15ceb8d3856effc8, points: 1, grade: true, solution: false -->
```python
assert(somme(13, -4) == 9)
```

Écrire une fonction `egaux` qui prend en argument deux entiers et teste s'ils sont égaux.

<!-- grade_id: cell-fb4382e7c5e63095, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-9b61c83ed53838ab, points: 1, grade: true, solution: false -->
```python
assert(egaux(2, 2) == True)
assert(egaux(2, 3) == False)
```

Écrire une fonction `surface` qui prend en argument un nombre et renvoie la surface du cercle de rayon ce nombre s'il est positif ou nul, et 0 sinon (prenez `3.14` comme valeur de $\pi$).

<!-- grade_id: cell-cb418dfaa47faace, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-6f57a7595a7416fd, points: 1, grade: true, solution: false -->
```python
assert(surface(-1) == 0)
assert(surface(0) == 0)
assert(surface(2) == 12.56)
```

Écrire une fonction `identite` qui prend en argument un entier et renvoie une chaîne de caractères "M." si l'entier vaut 1, "Mme" si l'entier vaut 2 et "Mx" sinon.

<!-- grade_id: cell-d0d20ea637201475, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-34a489ea3d6372aa, points: 1, grade: true, solution: false -->
```python
assert(identite(1) == "M.")
assert(identite(2) == "Mme")
assert(identite(7) == "Mx")
```

## Boucles et chaînes de caractères
Écrire une fonction `eternuement` qui prend en argument un entier `n` et affiche `ahahah...tchoum!!!` (`n` fois `ah` au début). Pour afficher `'toto'` sans retour à la ligne, on peut utiliser `print('toto', end='')`., pour afficher `'toto'` avec un retour à ligne, on utilise `print('toto')`.

<!-- grade_id: cell-2e40a97392cf764b, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

Vérifiez que vous avez l'affichage suivant :
    
```python
eternuement(5)
ahahahahah...tchoum!!!
```

```python
eternuement(5)
```

Écrire une fonction `epeler` qui prend en argument une chaîne de caractères et affiche ses caractères, un par ligne. (Rappel: on parcourt une chaîne de caractères comme une liste.)

<!-- grade_id: cell-22786418a677af66, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

Vérifiez que vous avez l'affichage suivant :

```python
epeler('abracadabra')
a
b
r
a
c
a
d
a
b
r
a
```

```python
epeler('abracadabra')
```

Écrire une fonction `aeiouy` qui prend en argument une chaîne de caractères et renvoie le nombre de voyelles qu'elle contient.

<!-- grade_id: cell-5ba70242356927eb, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-f794c6e056a6a4d0, points: 1, grade: true, solution: false -->
```python
assert(aeiouy('shfjdsfdsh') == 0)
assert(aeiouy('ahfdsjfhd') == 1)
assert(aeiouy('ghdfgfdafjsdkfd') == 1)
assert(aeiouy('ghdfhgfa') == 1)
assert(aeiouy('a') == 1)
assert(aeiouy('') == 0)
assert(aeiouy('vdwhfnvdkhvwiowfewfndsjfh') == 3)
```

(*) En Français, on peut couper les mots en fin de ligne avec un trait d'union, mais pas n'imorte où. Écrire une fonction `cesure_simplifiee` qui prend en argument une chaîne de caractères et la découpe en sous-chaînes (avec affichage de ces sous-chaînes) en suivant cet algorithme simplifié:
* on ne coupe pas avant la dernière lettre du mot,
* on ne coupe pas après la première lettre du mot si c'est une consonne,
* sinon:
    * on coupe entre deux consonnes,
    * on coupe après une voyelle si elle est suivie d'une consonne.

<!-- grade_id: cell-7d9e394e1f9f5237, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

Vérifiez que vous obtenez :
    
```python
cesure_simplifiee('proposition')
pro-
po-
si-
tion
```

et :

```python
cesure_simplifiee('vacances')
va-
ca-
n-
ces
```

```python
cesure_simplifiee('proposition')
```

```python
cesure_simplifiee('vacances')
```

(\**) Écrire une fonction `sapin` qui prend en argument un entier et affiche un sapin sous la forme suivante :

```
sapin(4)
   **
  *  *
 *    *
*      *
********
   ||

sapin(3)
  **
 *  *
*    *
******
  ||
```

<!-- grade_id: cell-5e33dce45b42bae4, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

Vérifiez vos sapins:

```python
sapin(3)
```

```python
sapin(4)
```

## Listes
Écrire une fonction `est_isocele` qui prend en argument trois listes supposées chacune représenter un point du plan (chaque liste contient deux nombres, l'abscisse et l'ordonnée du point) et teste si le triangle formé par ces trois points est isocèle ou pas. N'oubliez pas qu'on ne doit jamais tester l'égalité de deux nombres flottants.

<!-- grade_id: cell-32a277f38f3c0dc6, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-0734dd5b07851a32, points: 1, grade: true, solution: false -->
```python
assert(est_isocele([0, 0], [10, 0], [0, 10]) == True)
assert(est_isocele([2, 3], [6, 6], [-1, 7]) == True)
assert(est_isocele([4.3, -2.8], [5.1, -6.2], [0.9, -2]) == True)
assert(est_isocele([2, 6], [9.4, 1], [-1.2, 5]) == False)
```

Écrire une fonction `nb_vides` qui prend une liste de chaînes de caractères et renvoie le nombre de chaînes vides de cette liste.

<!-- grade_id: cell-f2486c2d7887d253, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-9332763caf14d35a, points: 1, grade: true, solution: false -->
```python
assert(nb_vides([]) == 0)
assert(nb_vides(['']) == 1)
assert(nb_vides(['toto']) == 0)
assert(nb_vides(['', 'toto', 'titi', '', '', 'lulu', '']) == 4)
```

Écrire une fonction `zip` qui prend en argument deux listes et renvoie une liste de même longueur que la plus courte liste passée en argument et dont la case `i` contient une liste de longueur 2 contenant les éléments des cases `i` des deux listes passées en argument.

exemple :

```python
zip([1, 2, 3, 4], ['a', 'b', 'c'])
[[1, 'a'], [2, 'b'], [3, 'c']]
```

<!-- grade_id: cell-f46186e1e3c8cb6b, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-e55420d64bfab7f4, points: 1, grade: true, solution: false -->
```python
assert(zip([], []) == [])
assert(zip([3, 2, 1], []) == [])
assert(zip([], ['toto', 'titi']) == [])
assert(zip([1, 2, 3, 4], ['a', 'b', 'c']) == [[1, 'a'], [2, 'b'], [3, 'c']])
```

Écrire une fonction `mesurer` qui prend en argument une liste de listes et renvoie une liste de même longueur dont la case `i` vaut la longueur de la liste dans la `i`e case de l'argument.

<!-- grade_id: cell-2b227f2285faee11, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-190a3d0696213b93, points: 1, grade: true, solution: false -->
```python
assert(mesurer([]) == [])
assert(mesurer([[]]) == [0])
assert(mesurer([[], [1, 2, 3], ['a'], [], [[1, 2, 3], [4, 5]]]) == [0, 3, 1, 0, 2])
```

(\**) Écrire une fonction `consecutifs` qui prend en argument une liste d'éléments et renvoie le nombre maximal d'éléments consécutifs identiques.

<!-- grade_id: cell-6a4cee303f542441, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-d6169ce616b7f4b0, points: 1, grade: true, solution: false -->
```python
assert(consecutifs([]) == 0)
assert(consecutifs(['', '', 'a']) == 2)
assert(consecutifs([1, 1, 2, 3, 3, 3, 3, 4, 4, 2, 2, 2]) == 4)
```

## Utilisation de modules
Pour faire cet exercice, vous devez installer le module `matplotlib` sur votre compte avec la commande suivante (c'est un peu long, c'est normal) :

```python
!pip install --user matplotlib  ##ceci n'est pas du python
```

Le module `matplotlib` contient un sous-module pyplot, importez-le de façon à pouvoir utiliser la commande qui suit (c'est un peu long la première fois, c'est normal):

<!-- grade_id: cell-ed12c83667748f71, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit

weights = [30, 50, 45, 80, 76, 55, 45, 47, 50, 65]
pyplot.hist(weights)
pyplot.show()
```

Pour dessiner des points, on peut donner en argument à la fonction `plot` du sous-module `pyplot` deux listes en argument, la première contenant les abscisses des points, la deuxième leurs ordonnées.

Écrire une fonction `representation` qui prend en argument une autre fonction et une liste de nombres et dessine les valeurs de la fonction passée en argument en les points d'abscisses fournies dans la liste.

exemple : le code
```python
import math

representation(math.cos, [.2*n for n in range(-40, 40)])
```
doit aboutir à l'affichage
![](plot_cos.png)

<!-- grade_id: cell-4a302edf954cb3f1, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

```python
Testez votre code :
```

```python
import math

representation(math.cos, [.2*n for n in range(-40, 40)])
```