# ITC - MPSI
---

# TP0 : apprendre à utiliser l'environnement de travail

Le but de cet énoncé est de faire en sorte que vous puissiez utiliser l'environnement de travail des TP de cette année.

Chaque semaine sera organisée de la manière suivante :
* \[au second semestre uniquement\] cours le mercredi matin
* TP le jeudi après-midi, à rendre _via le site_ à la fin du TP
* TPbis en devoir à la maison mis sur le site le vendredi soir au plus tard (lendemain du TP) et à rendre pour le mardi suivant au plus tard

Les corrections seront disponibles sur le site de TP.

Vous pouvez travailler sur n'importe quel ordinateur qui possède une connection internet et un navigateur web, par exemple au CDI. Pensez à sauvegarder en local au cas où.

Quand le traitement associé à une cellule est terminé (vous avez fini de la lire, vous avez fini d'écrire le code que vous aviez à écrire, vous voulez simplement exécuter le code sans le modifier), il faut utiliser la combinaison de touche `Shift-Enter` (ou `Maj-Entrée`): cela exécute ce qui doit être éventuellement exécuté, et passe à la cellule suivante.

---
## Quelques renseignements
Pour commencer, vous allez me donner quelques renseignements vous concernant, sous forme de chaînes de caractères `python`. En `python`, une chaîne de caractères peut prendre plusieurs formes :
* si elle est écrite sur une seule ligne, elle est encadrée de caractères `'` ou de caratères `"`;
* si elle est écrite sur plusieurs lignes (avec des retours à la ligne), elle est encadrée de caractères `'''` ou de caractères `"""`.

Un exemple :

```python
"Ceci est une chaîne sur une seule ligne"
'''Tandis que ceci est une chaîne
sur deux lignes.'''
```

Afin de me transmettre des renseignements vous concernant, il vous faudra donner un nom à chacune des chaînes de caractères contenant un de ces renseignements, c'est-à-dire créer une variables qui l'identifiera. Pour faire cela en `python`, il suffit d'écrire le nom de la variable, l'opérateur d'affectation `=`, puis la valeur qu'on souhaite associer à la variable.

Des exemples :
```python
pi = 3.14
perimetre_cercle_unitaire = 2*pi
prenom_Rousseau = "Jean-Jacques"
```

Créer une variable `nom` contenant votre nom.

<!-- grade_id: nom, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-122674bbcf04ad42, points: 1, grade: true, solution: false -->
```python
assert(len(nom)>0)
```

Créer une variable `prenom` contenant votre nom.

<!-- grade_id: prenom, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-9918e670dbcc6dfa, points: 1, grade: true, solution: false -->
```python
assert(len(prenom)>0)
```

```python
print("Bonjour, " + prenom )
```

---

## Valider et soumettre sa réponse
Si une erreur se produit à la validation d'une cellule par `Maj-Entrée`, elle sera affichée sur fond rose juste après la cellule, avec un message qui à terme vous aidera à identifier l'erreur.

Quand vous êtes sûr de votre rendu, vous pourrez tout retester grâce au bouton `Validate` en haut vers le milieu.

Enfin, pour rendre le devoir, vous devrez vous rendre à la page "Assignments" à partir de votre page d'accueil et cliquer sur `Submit`.

Créez maintenant une variable `info_points_positifs` qui contient une description (a priori sur plusieurs lignes, donc attention aux guillements utilisés) sur ce que vous aimez en informatique.

<!-- grade_id: positifs, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-8183883dc1c7a1e4, points: 1, grade: true, solution: false -->
```python
assert(len(info_points_positifs)>0)
```

Enfin, créez une variable `info_points_negatifs` qui contient une description de ce que vous n'aimez pas en informatique. (Ne vous censurez pas, ça me permettra d'adapter mes énoncés à vos attentes.)

<!-- grade_id: negatifs, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-649e0b3167cb211e, points: 1, grade: true, solution: false -->
```python
assert(len(info_points_negatifs)>0)
```

Créez une variable `mon_avenir` qui contient une description du ou des domaines dans lesquels vous auriez envie de travailler.

<!-- grade_id: cell-4ac5a6cdbdd5a38e, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-912e74effeca1018, points: 1, grade: true, solution: false -->
```python
assert(len(mon_avenir)>0)
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>

* on peut écrire des chaînes de caractères sur une ou plusieurs lignes en `python`; dans le cas de plusieurs lignes, il faut les encadrer par `'''` ou `"""`;
* une _variable_ est un nom qui permet d'identifier une valeur, l'action qui permet cette identification s'appelle un _affectation_, et on utilise l'opérateur `=` pour la faire;
* l'opérateur d'affectation `=` n'est pas symétrique : à gauche on met le nom de la variable, à droite la valeur qu'elle doit identifier.
    
    </div>