# ITC - MPSI
---

# TP1 : calcul de la racine carrée
## ou pourquoi il ne suffit pas de trouver un algorithme

Dans ce TP, nous allons étudier deux algorithmes de calcul de racine carrée, et comparer expérimentalement leur vitesse de convergence. Ce sera aussi l'occasion de revoir les bases de `python` que vous avez dû voir en spécialité Mathématiques au lycée.

## Rappel : écrire une fonction en `python`
Le but d'une fonction est d'écrire une boîte noire à laquelle on fournit des valeurs, et qui restitue un résultat (comme en mathématiques!).
En `python`, on définit une fonction de la manière suivante :

* mot-clef `def`
* nom de la fonction
* liste des paramètres entre parenthèses (_même s'il n'y a pas de paramètre_)
* deux points (`:`)
* retour à la ligne et indentation
* corps de la fonction (c'est-à-dire tout le code qui correspond aux calculs de la fonction)
* ligne vide (ou fin du fichier)

La valeur retournée par la fonction (c'est-à-dire la valeur qu'on attend d'elle) est précédée du mot-clef `return`, et une fois ce mot-clef rencontré dans un chemin d'exécution, on sort de la fonction pour retourner là où l'appel de fonction a été fait.

Il faut prendre l'habitude de documenter la fonction, par exemple par une chaîne de caractères juste avant.

Un exemple :
```python
import math

"""calcule l'aire d'un disque
r : rayon, nombre (entier ou flottant)
retour : aire du disque, nombre flottant"""
def aire_disque(r):
    return math.pi * r**2
```

(Un _flottant_ est la façon dont est représenté un nombre réel, nous reviendront là-dessus plus tard.)

Documenter et compléter la fonction suivante qui prend en argument deux nombres et renvoie leur moyenne arithmétique.

<!-- grade_id: cell-b28534382f6eb7aa, points: 0, grade: false, solution: true -->
```python
"""a, b : 2 nombres
resultat : moyenne arithmetique de a et b"""
def moyenne(a, b):
    return (a+b)/2
```

<!-- grade_id: cell-122674bbcf04ad42, points: 1, grade: true, solution: false -->
```python
epsilon = 0.00001
assert(abs(moyenne(0,0)) < epsilon)
assert(abs(moyenne(1,1) - 1) < epsilon)
assert(abs(moyenne(-1,1)) < epsilon)
assert(abs(moyenne(3,8) - 5.5) < epsilon)
```

## Tester l'égalité de deux flottants
On ne peut représenter tous les réels sur un ordinateur, car il y a une infinité de réels alors que la mémoire d'un ordinateur est finie. Pour représenter les réels, on utilise des _flottants_ (ou _nombres à virgule flottante_).
Pour des raisons que nous verrons plus tard dans l'année, les calculs sur flottants entraînent beaucoup d'erreurs d'arrondis, comme vous pouvez le constater avec le calcul suivant :

```python
0.3-0.2 == 0.1
```

Pour cette raison, il ne faut **jamais** tester si deux flottants sont égaux ou différents, mais tester si la valeur absolue de leur différence est suffisamment petite ou suffisamment grande:

```python
epsilon = 0.000001
abs((0.3-0.2) - 0.1) < epsilon
```

## Rappel : conditionnelle

Les structures de contrôle permettent de modifier le flot d'exécution du programme. En particulier `if` et `if ... else` introduisent des embranchements.

```python
if condition:
    code    #indentation obligatoire
```
se comporte de la manière suivante :
<img src="flot_if.png" width="200">


tandis que
```python
if condition:
    code1    #indentation obligatoire
else:
    code2    #indentation obligatoire
```
se comporte comme suit :
<img src="flot_if_else.png" width="230">

Documenter et compléter (sans utiliser la fonction native `abs` de `python`) la fonction suivante qui prend en argument un nombre et renvoie sa valeur absolue.

<!-- grade_id: cell-a81e45a52f0b7323, points: 0, grade: false, solution: true -->
```python
def valeur_absolue(x):
    if x >= 0 :
        return x
    else : 
        return -x
```

<!-- grade_id: cell-68cdf7728f1c9e85, points: 1, grade: true, solution: false -->
```python
epsilon = 0.00001
assert(abs( valeur_absolue(0) ) < epsilon)
assert(abs( valeur_absolue(10) - 10 ) < epsilon)
assert(abs( valeur_absolue(-15.3) - 15.3 ) < epsilon)
```

## Rappel : boucles conditionnelles

L'intérêt d'utiliser un programme pour résoudre un problème, réside dans le fait qu'on peut demander au programme de répéter beaucoup de fois une action (c'est plus rapide qu'à la main). Il existe pour cela plusieurs instructions en `python`, en particulier l'instruction `while`.

```python
while condition:
    code    #indentation obligatoire
```
se comporte de la manière suivante:
<img src="flot_while.png" width="230">

Si la condition est fausse dès le départ, alors le bloc d'instruction dans la boucle (`code`) n'est jamais exécuté.

## Racine carrée : calcul par dichotomie
Dans cette partie, nous allons aborder le problème du calcul de la racine carrée d'un nombre supposé supérieur ou égal à 1 (sans vérifier ce point), par une stratégie algorithmique classique, la _dichotomie_. Le principe de la dichotomie est le suivant :
* on cherche la réponse dans un intervalle,
* si la réponse est au milieu de l'intervalle, on a trouvé la solution,
* sinon
  * soit elle est inférieure à ce milieu et on regarde dans la moitié de gauche en suivant la même méthode,
  * soit elle est supérieure et on regarde dans la moitié de droite en suivant la même méthode.
C'est le principe que vous utilisez quand vous cherchez un mot dans un dictionnaire.

Ici, si on note $x$ le nombre dont on cherche la racine carrée, on sait que celle-ci appartient à l'intervalle
$$[1,x].$$

(Rappel : on suppose que $x\geq 1$.) On peut alors comparer le carré de la valeur candidate (ici $(x+1)/2$) à $x$ : s'ils sont suffisamment proches, on considère qu'on a trouvé $\sqrt{x}$, sinon on continue dans la bonne moitié de l'intervalle.

L'algorithme se traduit alors de la façon suivante:
* gauche $\leftarrow$ 1
* droite $\leftarrow$ x
* candidat $\leftarrow$ moyenne(gauche, droite)
* tant que $\text{candidat}^2$ trop éloigné de $x$
  * si $\text{candidat}^2 < x$, alors gauche $\leftarrow$ candidat &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;#car la racine se trouve dans $[\text{candidat},\text{droite}]$
  * sinon, droite $\leftarrow$ candidat &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;#car la racine se trouve dans $[\text{gauche},\text{candidat}]$
  * on met candidat à jour
* renvoyer la valeur de candidat &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;#si on sorti de la boucle, c'est que $\text{candidat}^2$ est suffisamment proche de $x$

Documenter et compléter la fonction suivante :

<!-- grade_id: cell-e9268bd0e190a646, points: 0, grade: false, solution: true -->
```python
def trop_eloigne(a, b, precision):
    return valeur_absolue(a-b) > precision

def racine_par_dichotomie(x, precision):
    gauche = 1
    droite = x
    candidat = (gauche+droite)/2
    while trop_eloigne(candidat**2, x, precision):
        if candidat**2 < x:
            gauche = candidat
        else : 
            droite = candidat
        candidat = (gauche+droite)/2
    return candidat
```

<!-- grade_id: cell-532fd3b19f9e2a48, points: 1, grade: true, solution: false -->
```python
import math
import random

epsilon = 0.0001
precisiontest = 0.00000001
assert(abs(racine_par_dichotomie(1, precisiontest) - 1) < epsilon)
assert(abs(racine_par_dichotomie(2, precisiontest) - math.sqrt(2)) < epsilon)
xtest = random.randint(43432,65364)
assert(abs(racine_par_dichotomie(xtest, precisiontest) - math.sqrt(xtest)) < epsilon)
```

Pour connaître le temps d'exécution d'une cellule, il suffit de mettre `%%time` en première ligne de la cellule :

```python
%%time
racine_par_dichotomie(xtest, precisiontest)
```

Output:
```
CPU times: user 32 µs, sys: 0 ns, total: 32 µs
Wall time: 35 µs
208.4202485364558
```

Comme sur un seul essai on peut tomber sur un cas particulier qui demande peu d'étapes ou au contraire en demande beaucoup, on va faire une boucle sur un ensemble de valeurs (ce type de boucle sera expliqué au TP2):
[ Attention : le calcul prend quelques secondes, soyez patients ]

```python
%%time
x = 43432
for i in range(100000):     # on fait 100000 essais
    racine_par_dichotomie(x+i, precisiontest)
```

Output:
```
CPU times: user 2.51 s, sys: 4.04 ms, total: 2.51 s
Wall time: 12.1 s
```

## Racine carrée : algorithme de Héron
Chez les Grecs anciens, trouver la racine carrée de $x$ était un problème de gémétrie : cela revenait à trouver un carré de côté $c$ dont l'aire valait $x$. L'algorithme de Héron repose sur cette vision : 
* on part d'un rectangle dont on sait que l'aire est $x$ : le rectangle $1\times x$,
* tant que le rectangle ne peut pas être considéré comme un carré (c'est-à-dire tant que ses côtés sont suffisamment différents)
  * on considère un nouveau rectangle dont un des côtés vaut la moyenne arithmétique des côtés de l'ancien rectangle, et l'autre ce qu'il faut pour que l'aire de ce rectangle soit encore $x$
* on retourne $x$.

Commenter et compléter la fonction suivante:

<!-- grade_id: cell-4534bcfab8b31593, points: 0, grade: false, solution: true -->
```python
def racine_Heron(x, precision):
    cote1 = 1
    cote2 = x
    while trop_eloigne(cote1, cote2, precision):
        cote1 = (cote1+cote2)/2
        cote2 = x/cote1
    return cote1
```

<!-- grade_id: cell-ebf553428a064e00, points: 1, grade: true, solution: false -->
```python
assert(abs(racine_Heron(1, precisiontest) - 1) < epsilon)
assert(abs(racine_Heron(2, precisiontest) - math.sqrt(2)) < epsilon)
assert(abs(racine_Heron(xtest, precisiontest) - math.sqrt(xtest)) < epsilon)
```

Vous pouvez comparer le temps d'exécution de ce nouvel algorithme avec le précédent.

```python
%%time
x = 43432
for i in range(100000):     # on fait 100000 essais
    racine_Heron(x+i, precisiontest)
```

Output:
```
CPU times: user 460 ms, sys: 27 µs, total: 460 ms
Wall time: 2.25 s
```

Faites le même test avec la fonction `sqrt` du paquet `math` pour vous convaincre qu'aucun des deux algorithmes n'est celui qui est utilisé dans les programmes informatiques (pour les curieux l'algorithme utilisé est une variante de l'algorithme de Cordic, qui repose sur des maths et de l'info que vous n'avez pas encore vues) :

[ au passage, notez comment on peut utiliser une fonction définie dans un paquet ]

```python
%%time
import math

x = 43432
for i in range(100000):     # on fait 100000 essais
    math.sqrt(x)
```

Output:
```
CPU times: user 16.7 ms, sys: 0 ns, total: 16.7 ms
Wall time: 83.6 ms
```

Rendez votre devoir en passant par la page "Assignments".

<div class="alert alert-success">
    <h2>Les points à retenir</h2>

* on ne teste jamais si deux flottants sont égaux ou pas : on teste s'ils sont suffisamment proches pour qu'on puisse les considérer comme égaux, ou suffisamment éloignés pour qu'on puisse les considérer comme différents;
* toute fonction doit être commentée par une chaîne de caractères juste avant qui
  * résume ce que fait la fonction
  * donne les types et rôles des paramètres
  * donne le type et la signification de la valeur de retour (s'il y en a une)
* principe de la dichotomie : on a un ensemble ordonné de valeurs candidates, on regarde si celle du milieu est la bonne, sinon on regarde dans la bonne moitié
* syntaxe et comportement de `if` et `if ... else`
* syntaxe et comportement de `while`
* pour utiliser la fonction `f` définie dans le paquet `paq`, on peut écrire `import paq` (au début du programme) puis invoquer `f` par `paq.f`
</div>