# ITC - MPSI
---

# TP10 : Compréhension de listes
On a déjà vu comment écrire une liste _in extenso_ (exemple : `[1,2,3]`) et comment ajouter un élément à une liste avec `append`. Nous explorons ici une autre méthode de construction de liste appelée _compréhension de liste_. Cette méthode consiste à partir d'une liste existante (ou d'une suite d'éléments donnée sous une autre forme, comme on le verra), et à déduire, à partir des éléments de cette liste, les éléments de la nouvelle liste.


Supposons qu'on dispose d'une liste d'entiers `L` et on veut obtenir une nouvelle liste qui contient les doubles des éléments de `L`. Avec ce que vous connaissez pour l'instant, vous devez écrire:

```python
L = [23, -1, 5, 92, -543, 32, 57] #exemple de liste pour les tests
doubleL = []
for i in L:
    doubleL.append(2*i)
doubleL
```

En fait, on peut faire nettement plus court:

```python
L = [23, -1, 5, 92, -543, 32, 57] #exemple de liste pour les tests
[ 2*i for i in L ]
```

Cette opération
```python
[ e(x) for x in L ]
```
s'appelle une _compréhension de liste_ (traduction horrible de l'anglais, mais c'est le terme effectivement utilisé).

Elle signifie qu'on parcourt les éléments de `L` (`for x in L`) et que pour chaque élément on évalue une certaine expression (`e(x)`). On obtient donc une liste de même longueur que `L` et telle que la case `i` de cette liste contient l'évaluation de l'expression `e(L[i])`.

On peut appliquer la même opération à tout type de `python` qui est _itérable_ (c'est-à-dire sur lequel on peut faire un `for`); on obtient à chaque fois une nouvelle liste.

Exemple:

```python
[ i**3 for i in range(5) ]   # donne [0, 1, 8, 27, 64]
[ i*2 for i in 'bonjour' ]   # donne ['bb', 'oo', 'nn', 'jj', 'oo', 'uu', 'rr']
```

### Exercice 1 : Passage en majuscule
La fonction `str.upper` renvoie une copie de son argument (une chaîne de caractères), en passant toutes les lettres en majuscules (il s'agit en fait de la fonction `upper` de la classe `str`, mais la notion de classe est hors programme).

Écrire et documenter une fonction `upper` qui prend en argument une liste de chaînes de caractères et renvoie une nouvelle liste avec une copie de ces chaînes en majuscules, en utilisant une compréhension de liste.

<!-- grade_id: cell-75e74172cad64643, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-028a8b1c5c280af8, points: 1, grade: true, solution: false -->
```python
assert(upper([]) == [])
assert(upper(['toto', '', '123', 'titi']) == ['TOTO', '', '123', 'TITI'])
```

### Exercice 2 : Nombre de caractères
Écrire et documenter une fonction `longueur` qui prend en argument une liste de chaînes de caractères et renvoie le nombre **total** de caractères dans l'ensemble de ces chaînes.

<!-- grade_id: cell-8696f3bb60bd00ca, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-7c266bcd28f0987d, points: 1, grade: true, solution: false -->
```python
assert(longueur([]) == 0)
assert(longueur(['toto', '', '123', 'titi']) == 11)
```

### Exercice 3 : Compréhensions imbriquées.

On peut imbriquer des compréhensions car l'expression calculée peut elle-même être calculée par compréhension, exemple:

```python
[ ['a'*i+'b'*j for j in range(i)] for i in range(10)]
```

Écrire et documenter une fonction `puissances` qui prend en argument une liste d'entiers `L` et un entier `n` et renvoie une liste de listes telle que la liste située en case `i` contienne les puissance `i`e des éléments de la liste d'origine, pour `i` allant de `0` à `n-1`. (Aucune contrainte sur la complexité, je veux juste que vous utilisiez des compréhensions de listes imbriquées.)

<!-- grade_id: cell-799933c2fc8c8aa5, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-9eff47d70a58d8b9, points: 1, grade: true, solution: false -->
```python
assert(puissances([], 5) == [[], [], [], [], []])
assert(puissances([-1], 5) == [[1], [-1], [1], [-1], [1]])
assert(puissances([-1, 7, 2, 3], 5) == [[1, 1, 1, 1],
 [-1, 7, 2, 3],
 [1, 49, 4, 9],
 [-1, 343, 8, 27],
 [1, 2401, 16, 81]])
```

### Exercice 4 : Intersection d'images
Dans cet exercice, nous représentons une image par une liste bidimensionnelle de booléens: `False` représente la couleur blanc, et `True` la couleur noir.

Un petit exemple avec une fonction d'affichage:

```python
def affiche(img):
    for ligne in img:
        for pixel in ligne:
            print(u'\u2B1B' if pixel else u'\u2B1C', end='')
        print()
                
affiche([[True, False, True, False], [True, True, False, True], [False, False, True, True]])
```

Écrire et documenter une fonction `intersection` qui prend en argument deux listes de listes de mêmes dimensions et rectangulaires contenant des `True` et des `False` et renvoie une liste de listes de mêmes dimensions dont une case vaut `True` si et seulement si les cases de mêmes coordonnées des listes fournies en argument valent toutes les deux `True`. Pas de boucle, uniquement des compréhensions de listes.

<!-- grade_id: cell-439ac1387fe249c6, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-2e6d1a7d303e228d, points: 2, grade: true, solution: false -->
```python
img1 = [[True, False, True, False], [True, True, False, True], [False, False, True, True]]
img2 = [ [True] * 4 for i in range(3)]
img3 = intersection(img1, img2)
assert(img3 == img1)

img2 = [ [False] * 4 for i in range(3)]
img3 = intersection(img1, img2)
assert(img3 == img2)

img2 = [ [True]*4, [False]*4, [True]*4 ]
img3 = intersection(img1, img2)
assert(img3 == [[True, False, True, False], [False, False, False, False], [False, False, True, True]])
```

### Exercice 4 : Liste de listes
On a parfois besoin de créer une liste de listes imbriquées d'une certaine taille, par exemple pour travailler sur une image. Une première idée (mauvaise) consiste à faire
```python
image = [ [ 0 ] * 10 ] * 15
```
Cette idée est mauvaise à cause de ce qui se passe en mémoire:
![](alloc.png)

On voit que la liste imbriquée (`[0] * 10`) est partagée par toutes les cases de `L`, ainsi on n'a pas vraiment 150 données, mais une suite de 10 données référencée 15 fois. La raison de cette situation est simple: on a vu que la valeur d'une liste n'est pas le contenu de ses cases, mais l'adresse où se trouve cette liste en mémoire, le `*15` externe recopie 15 fois cette valeur.

Il faudrait en fait fabriquer une nouvelle liste de 10 cases pour chaque case de `image`. Ceci peut se faire par compréhension de liste :
```python
image = [ [ 0 ] * 10 for i in range(15) ]
```
On se retrouve alors avec une situation tout à fait différente en mémoire :
![](alloc2.png)

On considère des représentations d'images au même format que dans l'exercice précédent (listes de listes de booléens). Écrire et documenter une fonction `mega_image` qui prend une représentation d'image dans ce format et  entiers `larg` et `haut` strictement positifs et fabrique la représentation d'une image qui correspond à la juxtaposition de l'image d'origine, à la manière d'Andy Warhol, mais sans modifier la couleur (dans cet exemple, les valeurs respectives de `larg` et `haut` seraient 5 et 3):
![](https://www.cnewyork.net/wp-content/uploads/resizefly/2018/11/warhol-marilyn-monroe-2-1200x750-696x435@1.jpg)

À cette fin, votre fonction commencera par créer une liste de listes aux bonnes dimensions, puis modifiera les contenus des cases.

<!-- grade_id: cell-e273d1d6e9327daf, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

```python
img = [[True, False, True, False], [True, True, False, True], [False, False, True, True]]
affiche(mega_image(img, 2, 2))
```

<!-- grade_id: cell-ec93bafdd6510267, points: 2, grade: true, solution: false -->
```python
img = [[True, False, True, False], [True, True, False, True], [False, False, True, True]]
assert(mega_image(img, 2, 3) == [[True, False, True, False, True, False, True, False],
 [True, True, False, True, True, True, False, True],
 [False, False, True, True, False, False, True, True],
 [True, False, True, False, True, False, True, False],
 [True, True, False, True, True, True, False, True],
 [False, False, True, True, False, False, True, True],
 [True, False, True, False, True, False, True, False],
 [True, True, False, True, True, True, False, True],
 [False, False, True, True, False, False, True, True]])
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* le principe de la compréhension de liste
* il n'est pas prudent de créer des listes multi-dimensionnelles en multipliant des listes de dimension inférieure
</div>