# ITC - MPSI
---

# TP11 : Algorithmes gloutons
Dans ce TP, nous allons explorer une stratégie algorithmique connue sous le nom d'_algorithme glouton_ (_greedy algorithms_ en anglais).

Le principe d'un algorithme glouton est de construire une solution au problème par étapes successives, en prenant à chaque étape la solution qui semble la meilleure, et en espérant que cette optimisation locale mènera à une solution globalement optimale. On parie donc sur les maxima locaux, en supposant que l'élément qui semble le meilleur aboutira au meilleur résultat final.

La qualité effective du résultat final dépendra en grande partie du choix de ce _meilleur_ élément. (En partie car tous les problèmes ne peuvent pas être résolus avec un algorithme glouton.)

Petit complément concernant les compréhensions de listes avant : on peut filtrer les données d'une collection en écrivant
```python
[ e(x) for x in L if f(x) ]
```
Par exemple si `L` est une liste d'entiers et qu'on veut prendre les moitiés des entiers paires de la liste, on peut écrire :

```python
L = [1, 2, 3, 4, 5, 9, 8, 7, -3, -2]  # pour l'exemple
[ x//2 for x in L if x%2==0 ]
```

## Le problème de l'emploi du temps

Considérons le problème d'attribution d'une salle pour un ensemble de cours :
* entrée : un ensemble de cours à placer, chaque cours possédant une heure de début et une heure de fin
* sortie : un ensemble maximal de cours qu'on peut placer sans chevauchement (par _maximal_ on entend qu'il n'existe pas d'ensemble de cours plaçables avec un cardinal strictement supérieur)

Le schéma général de résolution que nous allons adopter est le suivant :
* $S\leftarrow \emptyset$
* $P\leftarrow$ tous les cours avec accès rapide au meilleur choix
* tant que c'est possible
  * $e\leftarrow$ meilleur choix dans $P$
  * $S\leftarrow S\cup \{e\}$
  * mise à jour de $P$ (si besoin)

Pour simplifier le propos, on va supposer que chaque cours commence et termine à une heure entière. Un cours sera représenté par une liste de longueur 3 composée d'une chaîne de caractères donnant son nom, suivie de deux entiers positifs ou nuls représentant l'heure de début et l'heure de fin du cours.

exemples : `["math", 8, 12]` ou `["info", 14, 16]`

### Exo 1 : première tentative, au plus court
On peut penser que le mieux est de placer les cours qui durent le moins longtemps pour maximiser le nombre de cours qu'on place.

On connaît l'heure de début et l'heure de fin de chaque cours, il est donc très simple de placer les cours en fonction de leur durée.

1. Écrire et documenter une fonction `duree` qui prend en argument un triplet supposé représenter un cours et renvoie la durée de ce cours.

<!-- grade_id: cell-9854f500ad8e8ccc, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-b2a715c96a2cf411, points: 1, grade: true, solution: false -->
```python
assert(duree(["math", 8, 12]) == 4)
assert(duree(["hg", 0, 0]) == 0)
```

Je vous donne la fonction suivante (vous n'avez pas besoin de comprendre son code, mais elle ne fonctionnera correctement que si votre fonction `duree` est correcte):

```python
def tri_selon_critere(liste, critere):
    """entree : liste de cours, critere de tri
    sortie : la liste est triée selon le critere"""
    liste.sort(key = critere, reverse = True)


## exemple d'utilisation
L = [["math", 8, 12],  ["hg", 0, 0], ["info", 14, 16]]
tri_selon_critere(L, duree)
L
```

2. Écrire et documenter une fonction `mise_a_jour1` qui prend en argument une liste de cours, une heure de début `d` et une heure de fin `f`, et renvoie une nouvelle liste ne contenant que les cours de la liste fournie en paramètre qui terminent au plus tard à l'heure `d` ou commencent au plus tôt à l'heure `f`.

<!-- grade_id: cell-879e5e6c524fbb19, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

3. Écrire et documenter une fonction `emploi_du_temps1` qui prend en argument une liste de cours dans le format décrit précédemment, et renvoie une liste de cours qui ne se chevauchent pas, en suivant l'algorithme glouton, avec le choix du plus court comme étant le meilleur choix.

<!-- grade_id: cell-617c7f2ce25e08cd, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-eb71a23d72f4e94a, points: 2, grade: true, solution: false -->
```python
L = [["math", 8, 12],  ["hg", 0, 0], ["info", 14, 16]]
assert(sorted(emploi_du_temps1(L), key = lambda x : x[2]-x[1]) ==
      [['hg', 0, 0], ['info', 14, 16], ['math', 8, 12]])

L = [["math", 8, 12],  ["hg", 13, 14], ["info", 14, 16]]
assert(sorted(emploi_du_temps1(L), key = lambda x : x[2]-x[1]) ==
       [['hg', 13, 14], ['info', 14, 16], ['math', 8, 12]])

L = [["math", 8, 12],  ["hg", 13, 16], ["info", 14, 16]]
assert(sorted(emploi_du_temps1(L), key = lambda x : x[2]-x[1]) ==
       [['info', 14, 16], ['math', 8, 12]])
```

Le problème de cette solution en prenant d'abord le plus court, c'est qu'elle n'est pas optimale, voir par exemple l'entrée suivante qui place un seul cours, alors qu'elle pourrait en placer deux:

```python
L = [["math", 8, 12], ["phys", 12, 16], ["SI", 11, 13]]
emploi_du_temps1(L)
```

### Exo 2 : deuxième tentative, départ au plus tôt
On essaie maintenant une solution qui prend en priorité le cours qui commence le plus tôt parmi les cours encore compatibles avec ce qui a déjà été placé.

1. Écrire et documenter une fonction `debut` qui prend en argument un triplet supposé représenter un cours et renvoie l'heure de début de ce cours. Vous pourrez ensuite vous servir de cette fonction comme deuxième argument dans un appel à `tri_selon_critere`.

<!-- grade_id: cell-2ea929b53d0068c9, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-d7775ce144f34a65, points: 1, grade: true, solution: false -->
```python
assert(debut(['math', 9, 11]) == 9)
```

2. Écrire et documenter une fonction `mise_a_jour2` qui prend en argument une liste de cours et un entier `h` représentant une certaine heure, et renvoie une nouvelle liste ne contenant que les cours dont l'heure de début est au moins `h` (sans modifier l'ordre des cours). Utiliser une compréhension de liste pour construire une telle liste.

<!-- grade_id: cell-36700eb15ba9da77, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-0f759943bc8a2826, points: 2, grade: true, solution: false -->
```python
L = [["math", 8, 12],  ["hg", 0, 0], ["info", 14, 16]]
assert(mise_a_jour2(L, 4) == [['math', 8, 12], ['info', 14, 16]])

L = [["math", 8, 12],  ["hg", 0, 0], ["info", 14, 16]]
assert(mise_a_jour2(L, 13) == [['info', 14, 16]])

L = [["math", 8, 12],  ["hg", 0, 0], ["info", 14, 16]]
assert(mise_a_jour2(L, 15) == [])
```

3. Écrire et documenter une fonction `emploi_du_temps2` qui prend en argument une liste de cours dans le format décrie précédemment, et renvoie une liste de cours qui ne se chevauchent pas, en suivant l'algorithme glouton, avec le chiox au plus tôt comme étant le meilleur choix.

<!-- grade_id: cell-617c7f2ce25e08ce, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-3db677f9c8a3ccc8, points: 2, grade: true, solution: false -->
```python
L = [["math", 8, 12],  ["hg", 0, 0], ["info", 14, 16]]
assert(sorted(emploi_du_temps2(L), key = lambda x : x[1]) ==
       [['hg', 0, 0], ['math', 8, 12], ['info', 14, 16]])

L = [["math", 8, 12],  ["hg", 13, 14], ["info", 14, 16]]
assert(sorted(emploi_du_temps2(L), key = lambda x : x[1]) ==
       [['math', 8, 12], ['hg', 13, 14], ['info', 14, 16]])

L = [["math", 8, 12],  ["hg", 13, 16], ["info", 14, 16]]
assert(sorted(emploi_du_temps2(L), key = lambda x : x[1]) ==
       [['math', 8, 12], ['hg', 13, 16]])
```

Le problème de cette solution en prenant d'abord le cours qui commence le plus tôt, c'est qu'elle n'est pas optimale, voir par exemple l'entrée suivante qui place un seul cours, alors qu'elle pourrait en placer deux:

```python
L = [["math", 9, 12], ["phys", 12, 16], ["SI", 8, 13]]
emploi_du_temps2(L)
```

### Exo 3 : avec le moins d'incompatibilités
On pourrait penser qu'il faut commencer par prendre les cours qui présentent le moins d'incompatibilités, c'est-à-dire le moins d'intersection avec les autres cours.

1. Écrire et documenter une fonction `intersections` qui prend en argument une liste de cours `L` et un cours `C`, au format décrit précédemment, et renvoie le nombre d'intersections de `C` avec les cours de `L`. 

<!-- grade_id: cell-5072298a13436b32, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-81d3fd0ff299ccd9, points: 2, grade: true, solution: false -->
```python
L = [["math", 8, 12],  ["hg", 13, 16], ["info", 14, 16]]

assert(intersection(L, ["francais", 6, 7]) == 0)
assert(intersection(L, ["francais", 6, 9]) == 1)
assert(intersection(L, ["francais", 9, 10]) == 1)
assert(intersection(L, ["francais", 9, 14]) == 2)
```

2. Écrire et documenter une fonction `emploi_du_temps3` qui prend en argument une liste de cours dans le format décrit précédemment, et renvoie une liste de cours qui ne se chevauchent pas, en suivant l'algorithme glouton, avec le moins d'intrsection comme étant le meilleur choix (vous pouvez utiliser `mise_a_jour1`).

<!-- grade_id: cell-ba1cb8fbde828077, points: 0, grade: false, solution: true -->
```python
critere_intersection = lambda x: intersection(L, x)
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-6f6d59653c2a4ae5, points: 2, grade: true, solution: false -->
```python
L = [["math", 8, 12],  ["hg", 13, 16], ["info", 14, 16]]
assert(emploi_du_temps3(L) == [['info', 14, 16], ['math', 8, 12]])
```

3. Trouver un exemple d'entrée pour laquelle cet algorithme n'est pas optimal.

Describe the task here!

### Exo 4 : troisième méthode, fin au plus tôt
Montrer que si on choisit comme meilleur prochain choix, le cours qui se termine au plus tôt parmi les cours restants, alors on obtient un résultat optimal, c'est-à-dire qu'il n'existe pas d'ensemble de cours plaçables de cardinal supérieur à l'ensemble qu'on obtient par cette méthode.

Écrivez votre preuve dans la cellule suivante :

<!-- grade_id: cell-aa844a58656b80f7, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* le principe d'un algorithme glouton
* on ne trouve pas systématique une solution optimale avec un algorithme glouton
* compréhension de liste avec conditionnelle
* tri d'une liste avec `L.sort()` et `sorted(L)`, avec l'ordre naturel
</div>