# ITC - MPSI
---

# TP12 : Utilisation d'un module

Comme dans beaucoup de langages de programmation, l'installation de base en python contient certaines fonctionnalités, mais il existe des bibliothèques de fonctionnalités qu'on peut utiliser en le spécifiant à l'interpréteur. En `python`, ces bibliothèques s'appellent des _modules_.

Vous avez déjà utilisé un module. Comparez l'exécution des deux cellules suivantes (évidemment une fois la deuxième cellule exécutée, la première fonctionnera à son tours) :

```python
math.sqrt(8)
```

```python
import math
math.sqrt(8)
```

Il faut dont spécifier à l'interpréteur `python` qu'on veut utiliser un module avec le mot clef `import`. Ce mot-clef supporte plusieurs syntaxes:
```python
import math
```
et ensuite on peut utiliser tout ce qui apparaît dans le module `math` (fonction, variable globale ou autre) en écrivant par exemple `math.cos` ou `math.pi` (nom du module, suivi d'un point, suivi de l'identifiant de ce qui nous intéresse), ou alors
```python
from math import atan
```
et ensuite on peut utiliser l'identifiant directement, ici `atan` (sans rappeler le nom du module).

La deuxième solution permet de ne pas avoir à réécrire le nom du module, mais a l'inconvénient de ne pas autoriser l'utilisation de fonctions provenant de différents modules et qui auraient le même nom.

Pour avoir de la documenation sur une fonction, vous pouvez utiliser la fonction `help` avec en argument le nom de la fonction sur laquelle vous voulez de l'aide, par exemple :

```python
help(max)
```

Bien sûr, si cette fonction appartient à un module, il faudra d'abord faire l'import nécessaire puis identifier correctement la fonction.

`help` peut être utilisée pour sur propres fonctions :

```python
def f():
    """fonction inutile et qui ne prend aucun argument
    juste pour tester help"""
    pass

help(f)
```

Dans ce TP, nous allons utiliser un module appelé `ColabTurtlePlus`. Vous allez d'abord l'installer sur votre compte en exécutant la cellule suivante (une fois l'installation faite, le module reste installé, ce n'est pas la peine d'exécuter cette cellule à chaque fois, mais la première fois il faut relancer le noyau une fois l'installation faite) :

```python
pip install --user ColabTurtlePlus  ## ceci n'est pas du python
```

Vous devez ensuite importer la classe `Turtle` du module `ColabTurtlePlus`:

<!-- grade_id: cell-f474fe7ae80ec0d0, points: 0, grade: false, solution: true -->
```python
## écrivez ici l'import
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-20e0c943ddb70d48, points: 1, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

Les fonctions d'une classe s'utilisent avec la même syntaxe que les fonctions d'un module: nom de la classe, suivi d'un point, suivi du nom de la fonction. Voici un premier exemple de programme de dessin :

```python
Turtle.clearscreen()  ## à mettre au début de chaque nouveau dessin
Turtle.delay(0)       ## pour que le dessin commence tout de suite
Turtle.speed(10)      ## vitesse : 0 = à l'arrêt, 13 = le plus vite
Turtle.setup(300,300) ## dimensions de la fenêtre
Turtle.showborder()
for i in range(4):
  Turtle.forward(100)
  Turtle.left(90)
```

### Exercice 1 : Marche aléatoire d'une tortue ivre

Pour afficher une tortue plutôt qu'une flèche, il suffit d'utiliser l'appel `Turtle.shape('turtle')`. Écrire et documenter une fonction `marche_aleatoire` qui prend en argument un nombre de pas `n` et fait avancer la tortue de manière erratique : après chaque pas (un pas correspondant à une avancée de 20 unités), vous tirez au hasard l'angle de rotation de votre tortue, sans qu'elle puisse tourner de plus de 90 degré (une tortue ivre ne fait pas de gestes brusques) (fonction `randint` du module `random`).

<!-- grade_id: cell-bc19071532b16a8f, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

Vous pouvez maintenant regarder la marche aléatoire de votre tortue :

```python
marche_aleatoire(100)
```

### Exercice 2 : Trois chiens qui se poursuivent
Dans cet exercice, on considère trois chiens, chacun au sommet d'un triangle équilatéral, et chacun poursuivant le chien sur sa droite. On veut représenter les trajectoires de ces chiens.

1. À chaque unité de temps, chaque chien avance d'une unité en direction du chien qu'il poursuit. Écrire et documenter une fonction `avance` qui prend en argument deux listes de deux nombres: la première représente les coordonnées actuelles d'un chien, la deuxième les coordonnées de sa cible (supposées différentes, sans avoir à le vérifier). Cette fonction renvoie les nouvelle coordonnées du chien après une unité de temps, sous la forme d'une liste de deux nombres (rappel : la racine carrée est la fonction `sqrt` du module `math`).

<!-- grade_id: cell-4c1d9b14fcde49ae, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-91088165340166c2, points: 2, grade: true, solution: false -->
```python
import math

epsilon = 0.0000001

x, y = avance([0, 0], [1, 0])
assert(math.fabs(x-1)<epsilon and math.fabs(y)<epsilon)

x, y = avance([0, 0], [0, 1])
assert(math.fabs(x)<epsilon and math.fabs(y-1)<epsilon)

x, y = avance([0, 0], [1, 1])
assert(math.fabs(x-0.7071067811865475)<epsilon and math.fabs(y-0.7071067811865475)<epsilon)

x, y = avance([23, 89], [-12, 112])
assert(math.fabs(x-22.164294520298725)<epsilon and math.fabs(y-89.54917788666084)<epsilon)
```

2. Écrire et documenter une fonction `poursuite_infernale` qui prend en argument une liste de 3 sous-listes représentant les coordonnées initiales des trois chiens et un entier représentant le nombre de pas de la simulation et renvoie une liste contenant 3 listes, chacune composées des coordonnées successives d'un des chiens.

<!-- grade_id: cell-7efc97d8f0fb677b, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

3. Pour faire l'animation, je vous fournis une fonction qui prend en argument le résultat de `poursuite_infernale` et effectue l'affichage correspondant. Vous pouvez tester vos fonctions en dessous.

```python
def trace(trajectoires):
    Turtle.clearscreen()
    Turtle.setup(400,400)
    Turtle.showborder()
    chiens = [Turtle.Turtle(), Turtle.Turtle(), Turtle.Turtle()]
    for i in range(3):
        chiens[i].shape('blank')
        chiens[i].delay(0)
        chiens[i].speed(13)
        chiens[i].jumpto(tuple(trajectoires[i][0]))
                  
    for k in range(1, len(trajectoires[0])):
        for i in range(3):
            chiens[i].goto(trajectoires[i][k][0], trajectoires[i][k][1])
```

```python
r = 200
coord_initiales = [[r*math.cos(0), r*math.sin(0)], 
                   [r*math.cos(2*math.pi/3), r*math.sin(2*math.pi/3)],
                   [r*math.cos(4*math.pi/3), r*math.sin(4*math.pi/3)]]
trace(poursuite_infernale(coord_initiales, (int)(1.2*r)))
```

### Exercice 3 : Courbe du dragon

![](dragon10.png)

La _courbe du dragon_ peut être représentée par tableau de booléens. Elle est obtenue récursivement de la manière suivante (dans ce qui suit, je mets la suite des valeurs des cases du tableau, avec `F` pour `False` et `T` pour `True`) :

* $C_0$ : `F`
* $C_{n+1}$ : $C_n$ `F` $\tilde{C_n}$, où $\tilde{C}$ est construit à partir de $C$, en prenant le miroir du tableau (=cases lues de droite à gauche) et en inversant `True` et `False`.

$\def\F{\text{F}\,}
\def\T{\text{T}\,}$
exemple de passage de $C$ à $\tilde{C}$ :
$$\underbrace{\F\F\F\T\T\F}_{C} \xrightarrow{\text{miroir}} \F\T\T\F\F\F
\xrightarrow{\text{inversion}} \underbrace{\T\F\F\T\T\T}_{\tilde{C\ }}$$

premières valeurs de la courbe du dragon :

$$\begin{array}{lcl}
C_1 &=& \underbrace{\F}_{C_0} \F \underbrace{\T}_{\tilde{C_0}}\\
C_2 &=& \underbrace{\F\F\T}_{C_1} \F \underbrace{\F\T\T}_{\tilde{C_1}}\\
C_3 &=& \underbrace{\F\F\T\F\F\T\T}_{C_2} \F \underbrace{\F\F\T\T\F\T\T}_{\tilde{C_2}}
\end{array}$$

1. Écrire et documenter une fonction `mue` qui prend en argument un tableau de booléens et construit un nouveau tableau en utilisant la transformation décrite au-dessus (le tableau passé en paramètre ne doit pas être modifié). Vous pouvez utiliser :

* la fonction `reversed` qui permet d'inverser les éléments d'une liste,
* la méthode `extend` qui s'utilise comme `append` mais permet d'étendre une liste avec les éléments d'une autre liste, ainsi après la séquence :

```python
L = [1, 2, 3]
L.extend([9, 8, 7])
```

`L` vaut `[1, 2, 3, 9, 8, 7]`.

<!-- grade_id: cell-cb0c073dd359a3a2, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-fd314902c7df1e25, points: 2, grade: true, solution: false -->
```python
dragon = [ False ]
assert(mue(dragon) == [False, False, True] and dragon == [False])

assert(mue(mue(dragon)) == [False, False, True, False, False, True, True] and dragon == [False])
```

2. Écrire et documenter une fonction `courbe` qui calcule la $n$-ème courbe du dragon.

<!-- grade_id: cell-5c152af8e73bbffb, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-4306e044694ff0d3, points: 2, grade: true, solution: false -->
```python
assert(courbe(0) == [False])
assert(courbe(1) == [False, False, True])
assert(courbe(5) == [False, False, True, False, False, True, True, False, False, False, True, True, False, True, True, False, False, False, True, False, False, True, True, True, False, False, True, True, False, True, True, False, False, False, True, False, False, True, True, False, False, False, True, True, False, True, True, True, False, False, True, False, False, True, True, True, False, False, True, True, False, True, True])
```

3. Pour dessiner la courbe du dragon $C_n$, on part d'un segment unitaire ($((0,0), (u,0))$ par exemple) et on parcourt le tableau obtenu précédemment : à chaque `False`, on tourne à droite (de 90°) et à chaque `True` on tourne à gauche (de 90°), et on avance d'une unité. (La courbe présentée en début d'exercice est $C_{10}$.)

<!-- grade_id: cell-f2a6393815628455, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-1155b873f96a29a9, points: 2, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

```python
dessin(courbe(8), 10)
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* comment importer un module (`import module`) puis utiliser une fonction du module (`module.fonction`)
* comment importer une fonction spécifique d'un module (`from module import fonction`)
* `extend` pour étendre une liste (et comment s'en servir)
* utilisation de `help` pour avoir de la documentation sur une fonction
</div>