# TP13 : Fichiers
Dans ce TP nous allons apprendre à lire et écrire"dans un fichier qui se trouve sur le disque.

Quand un programme doit accéder au contenu d'un fichier, le système vérifie que cette lecture est autorisée, puis il met en place toute une structure de données pour permettre cette lecture. L'ensemble de ces opérations est désigné par le terme _ouverture d'un fichier en lecture ou en écriture_. Une fois ces opérations faites, le programme peut accéder au contenu du fichier. Quand il a terminé ses manipulations (lecture et/ou écriture), il doit signaler qu'il n'a plus besoin de ces ressources, afin qu'elles puissent être mises à jour si besoin sur le disque.

L'ouverture en lecture et/ou en écriture d'un fichier qui s'appelle `toto` se fait en `python` grâce à la fonction `open` à laquelle on fournit comme argument une chaîne de caractères qui contient le nom du fichier. Cette fonction renvoie un flux. Nous ne rentrerons pas dans les détails de ce qu'est un flux, le seul point important est que nous pourrons utiliser cette valeur de retour pour effectuer des lectures et des écritures dans le fichier ainsi ouvert.

```python
f = open("toto")
```

Output:
```
FileNotFoundError                         Traceback (most recent call last)
/tmp/ipykernel_14234/1239679143.py in <module>
----> 1 f = open("toto")

FileNotFoundError: [Errno 2] No such file or directory: 'toto'
```

Dans l'exemple précédent, le message d'erreur vient du fait qu'il n'y a aucun fichier qui s'appelle `toto`. Recommençons avec le nom d'un fichier existant (vous pouvez voir le contenu de ce fichier : [fichier_0](Fichiers/fichier_0)) :

```python
f = open('Fichiers/fichier_0')
```

Quand on termine d'utiliser les ressources associées au fichier, il faut les libérer, ce qui se fait par la méthode `close` (attention à la syntaxe) :

```python
f.close()
```

# Lecture d'un fichier
On peut lire d'un coup toutes les lignes d'un fichier (pas recommandé si le fichier est long), avec la méthode `read` :

```python
f = open('Fichiers/fichier_0')
f.read()
```

Output:
```
'Ceci est un exemple de fichier\nà lire\nsur plusieurs lignes.\n'
```

Vous remarquez que les retours à ligne apparaissent sous la forme du caractère `'\n'`. Si on affiche cette chaîne à l'aide de la fonction `print`, on a effectivement des retours à la ligne.

La pointe de lecture se déplace de gauche à droite au fur et à mesure, une fois la lecture de tout le fichier faite, il n'y a donc plus rien à lire :

```python
f.read()
```

Output:
```
ValueError                                Traceback (most recent call last)
/tmp/ipykernel_14580/3059900045.py in <module>
----> 1 f.read()

ValueError: I/O operation on closed file.
```

```python
f.close()
```

On peut également obtenir toutes les lignes d'un fichier sous la forme d'une liste de lignes (la coupure se fait alors à chaque changement de ligne) grâce à la méthode `readlines` :

```python
f = open('Fichiers/fichier_0')
lignes = f.readlines()
f.close()

lignes
```

Output:
```
['Ceci est un exemple de fichier\n', 'à lire\n', 'sur plusieurs lignes.\n']
```

On peut aussi demander lire ligne à ligne :

```python
f = open('Fichiers/fichier_0')
ligne1 = f.readline()
ligne2 = f.readline()
f.close()

ligne2
```

Output:
```
'à lire\n'
```

`readline` renvoie une chaîne vide (`''`) en fin de fichier.

```python
f = open('Fichiers/fichier_0')
f.read()
```

Output:
```
'Ceci est un exemple de fichier\nà lire\nsur plusieurs lignes.\n'
```

## Exercice 1 : compter les lignes
Écrire et documenter une fonction `wcl` qui prend en argument un nom de fichier et renvoie le nombre de lignes de ce fichier. Ne pas oublier de fermer le fichier.

<!-- grade_id: cell-92b9a462c3b9957e, points: 0, grade: false, solution: true -->
```python
###BEGIN SOLUTION
def wcl(fichier):
    f = open(fichier)
    res = len(f.readlines())
    f.close()
    return res
###END SOLUTION
```

<!-- grade_id: cell-fbd6efee6b39eb22, points: 2, grade: true, solution: false -->
```python
assert(wcl('Fichiers/fichier_1') == 0)
assert(wcl('Fichiers/fichier_2') == 1)
assert(wcl('Fichiers/fichier_3') == 3)
assert(wcl('Fichiers/fichier_4') == 10)
assert(wcl('Fichiers/fichier_5') == 10)
assert(wcl('Fichiers/fichier_6') == 2862)
```

## Exercice 2 : numéroter les lignes
Le but de cet exercice est d'écrire une fonction qui affiche les lignes d'un fichier les unes après les autres, en les numérotant. Pour cela, il faut pouvoir afficher du texte: on utilise la fonction `print` avec comme argument la chaîne de caractères à afficher :
```python
print('texte à afficher')
```

Quand on veut afficher un entier, on peut donner directement l'entier à `print`, qui se charge de la conversion :
```python
print(5)
n = 3
print(n)
```

En revanche si on veut combiner un entier et une chaîne, il faut explicitement convertir l'entier en utilisant la fonction `str` :
```python
n = 3
print('n vaut ' + str(n))
```

Écrire et documenter une fonction `catn` qui prend en argument un nom de fichier et affiche son contenu avec le numéro de chaque ligne en début de ligne.

<!-- grade_id: cell-fda9fe42b1180eb6, points: 2, grade: true, solution: true -->
```python
###BEGIN SOLUTION
def catn(fichier):
    f = open(fichier)
    i = 1
    for ligne in f.readlines():
        print(str(i) + ' ' + ligne)
        i = i + 1
###END SOLUTION
```

À vous de tester visuellement votre réponse. Voici les sorties attendues sur deux exemples:
```
catn('Fichiers/fichier_1')
```

```
catn('Fichiers/fichier_7')
1 ligne 1

2 ligne 2

3 

4 ligne 4
```


```python
catn('Fichiers/fichier_1')
```

```python
catn('Fichiers/fichier_7')
```

Output:
```
1 ligne 1

2 ligne 2

3 

4 ligne 4
```

## Exercice 3 : lignes vides
Écrire et documenter une fonction `lignes_vides` qui prend en argument un nom de fichier et renvoie le nombre de lignes vides qu'il contient (c'est-à-dire le nombre de lignes qui contiennent uniquement le caractère `'\n'`).

<!-- grade_id: cell-737a8df8b57f9696, points: 0, grade: false, solution: true -->
```python
###BEGIN SOLUTION
def lignes_vides(fichier):
    f = open(fichier)
    res = len([ ligne for ligne in f.readlines() if ligne=='\n' ])
    f.close()
    return res
###END SOLUTION
```

<!-- grade_id: cell-dc98bdefaba9fd7a, points: 2, grade: true, solution: false -->
```python
assert(lignes_vides('Fichiers/fichier_1') == 0)
assert(lignes_vides('Fichiers/fichier_2') == 1)
assert(lignes_vides('Fichiers/fichier_3') == 0)
assert(lignes_vides('Fichiers/fichier_4') == 7)
assert(lignes_vides('Fichiers/fichier_5') == 3)
assert(lignes_vides('Fichiers/fichier_6') == 113)
```

## Exercice  : Tester si deux fichiers sont identiques
Écrire et documenter une fonction `diff` qui prend en argument les noms de deux fichiers et teste si les deux fichiers sont identiques.

<!-- grade_id: cell-34082345e1699b69, points: 0, grade: false, solution: true -->
```python
###BEGIN SOLUTION
def diff(fic1, fic2):
    f1 = open(fic1)
    f2 = open(fic2)
    l1 = f1.readlines()
    l2 = f2.readlines()
    return l1 == l2
##END SOLUTION
```

<!-- grade_id: cell-6e4695b857b34c55, points: 2, grade: true, solution: false -->
```python
assert(diff('Fichiers/fichier_6', 'Fichiers/fichier_6') == True)
assert(diff('Fichiers/fichier_6', 'Fichiers/fichier_8') == True)
assert(diff('Fichiers/fichier_6', 'Fichiers/fichier_7') == False)
assert(diff('Fichiers/fichier_6', 'Fichiers/fichier_1') == False)
```

# Écriture dans un fichier
Pour écrire dans un fichier, il faut spécifier au moment de l'ouverture (`open`) qu'on veut écrire dans le fichier. Pour cela, on donne un deuxième argument à `open` sous forme d'une chaîne de caractères de longueur 1 qui explicite le type d'écriture:
* `'a'` pour écrire à la fin du fichier (a pour _append_)
* `'w'` pour écrire à la place du contenu du fichier qui est alors perdu (w pour _write_)

Pour écrire dans le fichier, on utilise alors la méthode `write` :
```python
f = open(fichier, 'w')
f.write('le texte à écrire')
f.close()
```

Contrairement à `print`, la méthode `write` n'ajoute pas de retour à la ligne en fin de chaîne. Si on veut un retour à la ligne, il faut l'ajouter explicitement, avec le caractère `'\n'`:
```python
f = open(fichier, 'w')
f.write('un texte avec retour à la ligne à la fin\n')
f.close()
```

Contrairement à `print`, la méthode `write` ne convertit pas toute seule son argument en chaîne de caractères.

## Exercice 4 : Écrire des nombres
Écrire et documenter une fonction `denombrer` qui prend en argument un entier positif ou nul et le nom d'un fichier et écrit dans ce fichier tous les entiers compris entre 0 et le premier argument (compris), avec un retour à la ligne après chaque entier.

<!-- grade_id: cell-5013620503297331, points: 0, grade: false, solution: true -->
```python
###BEGiN SOLUTION
def denombrer(n, fic):
    f = open(fic, 'w')
    for i in range(n):
        f.write(str(i) + '\n')
    f.close()
###END SOLUTION
```

<!-- grade_id: cell-65814ca4b8ecc2cf, points: 2, grade: true, solution: false -->
```python
fic = 'Fichiers/test1'
n = 5
denombrer(n, fic)
with open(fic) as f:
    assert(f.readlines() == [ str(i)+'\n' for i in range(n) ])
    f.close()
    
!rm Fichiers/test1  #pas du python
```

## Exercice 5 : Copier un fichier
Écrire et documenter une fonction `copie` qui prend en argument deux noms de fichiers et copie le contenu du premier dans le deuxième (en écrasant le contenu existant éventuel).

<!-- grade_id: cell-541384f63490d26e, points: 0, grade: false, solution: true -->
```python
###BEGIN SOLUTION
def copie(src, dest):
    s = open(src)
    d = open(dest, 'w')
    for ligne in s.read():
        d.write(ligne)
    s.close()
    d.close()
###END SOLUTION
```

<!-- grade_id: cell-510231615b1ad552, points: 2, grade: true, solution: false -->
```python
copie('Fichiers/fichier_3', 'Fichiers/copie')
assert(diff('Fichirs/fichier_3', 'Fichiers/copie') == True)
```