# ITC - MPSI
---

# TP15 : 🚗 🚙 🚕
Dans ce TP, nous allons récupérer et traiter des données réelles concernant les voitures commercialisées en France en mars 2014 (données récupérées sur https://www.data.gouv.fr/fr/datasets/emissions-de-co2-et-de-polluants-des-vehicules-commercialises-en-france/, j'ai juste supprimé la première ligne qui donne les significations des colonnes pour simplifier le traitement).

## Format `csv`
Le format `csv` est un des formats utilisés par les tableurs. C'est un format texte très simple:
* une ligne du tableau, c'est une ligne dans le fichier `csv`,
* dans chaque ligne, les contenus des cases apparaissent dans l'ordre (sans formules, uniquement des valeurs), séparés par un caractère appelé séparateur de champs, et éventuellement encadrés par des caractères appelés séparateurs de chaînes de caractères.

Dans ce TP, le séparateur de champs est le caractère `;` et il n'y a pas de séparateur de chaînes de caractères.

Pour découper une chaîne en fonction d'un séparateur, on peut utiliser la méthode `split`:

```python
'ceci est un exemple'.split(' ')   ## le séparateur est le caractère espace
```

```python
'ceci est un exemple'.split('e')   ## le séparateur est le caractère 'e'
```

```python
'ceci est un exemple'.split(' e')  ## le séparateur est la chaîne de caractères ' e'
```

Si on connaît à l'avance le nombre de morceaux obtenus, on peut dépaqueter directement:

```python
debut, fin = 'aaaaa-bbb'.split('-')
print('debut : ' + debut)
print('fin : ' + fin)
```

## Exercice 1 : les marques
Le fichier 'voitures.csv' donne des renseignements sur les voitures commercialisées en France en mars 2014.
La première colonne contient les marques des voitures.

1. Écrire et documenter une fonction `marques` qui prend en argument le nom d'un fichier supposé au format `csv` et renvoie la liste des valeurs de la première colonne.

<!-- grade_id: cell-41ff53652926bef6, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-f82f1e8bf48b2a51, points: 2, grade: true, solution: false -->
```python
marq = marques('voitures.csv')
assert(len(marq) == 55044)
assert(marq[4322] == 'MERCEDES')
assert(marq[51017] == 'VOLKSWAGEN')
```

2. Écrire une fonction `uniq` qui prend en argument une liste d'éléments comparables supposée triée et renvoie une liste avec ces éléments sans répétition. Cette fonction ne doit parcourir qu'une seule fois la liste passée en argument.

<!-- grade_id: cell-f276d82049923c94, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-b28bc41e968dd78a, points: 2, grade: true, solution: false -->
```python
marq = uniq(marques('voitures.csv'))
assert(len(marq) == 46)
assert(marq[26] == 'MERCEDES')
assert(marq[44] == 'VOLKSWAGEN')
```

## Exercice 2 : les modèles
On veut maintenant obtenir tous les modèles d'une marque donnée, toujours à partir du fichier 'voitures.csv'. Dans ce fichier, le nom d'un modèle apparaît dans la deuxième colonne, juste après la marque.

1. Écrire et documenter une fonction `modeles` qui prend en argument un nom de fichier supposé au format `csv` et avec le format précédemment décrit et renvoie un dictionnaire où les clefs sont des marques de voitures et les valeurs les listes de modèles de la marque, sans doublon (le fichier `csv` fourni est supposé être trié par marque puis pas modèle). Votre fonction ne doit parcourir les données qu'une fois.

<!-- grade_id: cell-ac75274ceaa7cfd4, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-04d48a69b4e8022b, points: 2, grade: true, solution: false -->
```python
mm = modeles('voitures.csv')
assert(len(mm['MERCEDES']) == 23)
assert(len(mm['VOLKSWAGEN']) == 21)
```

2. Écrire et documenter une fonction `la_plus_variee` qui prend en argument le nom d'un fichier au format `csv` comme décrit précédemment et renvoie le nom de la marque qui possède le plus de modèles différents.

<!-- grade_id: cell-9de817f2de610360, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-28d31a59e08db27d, points: 2, grade: true, solution: false -->
```python
assert(la_plus_variee('voitures.csv') == 'AUDI')
```

On a vu qu'on peut construire une liste à partir d'une autre liste grâce à une compréhension de liste. On peut utiliser le même genre d'opérations pour construire un nouveau dictionnaire.

3. Écrire et documenter une fonction `nb_de_modeles` qui prend en argument un dictionnaire tel que renvoyé par la fonction `modeles` et renvoie un nouveau dictionnaire où les clefs sont des marques de voitures et les valeurs le nombre de modeles associes.

<!-- grade_id: cell-3e20ded391e9bfbd, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-7ee46910b8f8c5b2, points: 2, grade: true, solution: false -->
```python
N = nb_de_modeles(modeles('voitures.csv'))
assert(N['HYUNDAI'] == 8)
assert(N['PORSCHE'] == 27)
```

## Exercice 3 : émission de $CO_2$ en fonction du type de moteur
Dans le fichier 'voitures.csv', la colonne n°6 (en commençant la numéroration à 0) fournit le type de carburant utilisé et la colonne n°14 l'émission de $CO_2$.

1. Écrire et documenter une fonction `carburants` qui renvoie la liste des divers types de carburant sans doublon (pensez à la fonction `sorted` que nous avons déjà rencontrée).


<!-- grade_id: cell-179b18c212aa0917, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-2267b5de5bdb80fc, points: 2, grade: true, solution: false -->
```python
assert(len(carburants('voitures.csv')) == 13)
```

2. Écrire et documenter une fonction `emission` qui renvoie un dictionnaire dont les clefs sont les différents types de carburant utilisés et les valeur la moyenne de $CO_2$ émis pour ce type de carburant. Vous pouvez utiliser les fonctions:
* `int` qui prend une chaîne de caractères supposée représenter un entier en argument et renvoie l'entier correspondant,
* `sum` qui avec une liste en argument renvoie la somme des éléments de la liste. 

<!-- grade_id: cell-f906ce341fa511c1, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-eaaae3d5f017a85f, points: 2, grade: true, solution: false -->
```python
e = emission('voitures.csv')
assert(abs(e['FE'] - 191.1) < 0.00001)
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* méthode `split` sur une chaîne de caractères
* fonction `sum` sur une liste
* fonction `int` pour transformer une chaîne en entier
* compréhensions pour créer un nouveau dictionnaire
</div>