# ITC - MPSI
---

# TP16 : Tranches

Dans ce TP, nous allons étudier une notation pour récupérer une partie contiguë (c'est-à-dire une _tranche_, _slice_ en anglais) d'une chaîne ou d'une liste (notation qui fonctionne également pour d'autres types que nous n'avons pas encore vus pour lesquels il y a un notion d'ordre des éléments, comme les tuples).

On dispose d'une liste

```python
L = ['zero', 'un', 'deux', 'trois', 'quatre', 'cinq', 'six', 'sept']
```

et nous voulons obtenir une liste qui contient les cases 2 à 5 de cette liste. Bien sûr nous pouvons créer une liste vide et faire un certain nombre d'appels à `append`, ou utiliser une compréhension de liste, mais `python` fournit une notation pour éviter d'écrire autant de code :

```python
L[2:6]
```

Cette notation `L[i:j]` fournit une _nouvelle_ liste dont les éléments sont, dans l'ordre, les contenus des cases `i` à `j-1` de `L` (attention : le contenu de la case `i` est inclus, celui de la case `j` est exclu).

La liste obtenue est bien une nouvelle liste :

```python
L = ['zero', 'un', 'deux', 'trois', 'quatre', 'cinq', 'six', 'sept']
M = L[2:6]
M[0] = '2'
print('M = ' + str(M))
print('L = ' + str(L))
```

Si on veut une tranche qui commence au début ou qui termine à la fin, on n'est pas obligé de mettre l'indicec correspondant (mais il faut laisser les `:`, sinon on n'accède qu'à une seule case):

```python
L[:6]
```

```python
L[2:]
```

Cela permet de faire une copie d'une liste à moindre effort:

```python
L[:]
```

Si on utilise des indices qui sortent des limites de la liste, seules les cases qui existent sont fournies:

```python
L[2:15]
```

Si on utilise des indices positifs avec un indice à gauche supérieur à l'indice à droite, on obtient une tranche vide:

```python
L[6:2]
```

On peut utiliser des indices négatifs qui désignent alors une position à partir de la fin de la liste (comme lorsqu'on accède à une case à partir d'un indice négatif):

```python
L[2:-1]
```

```python
L[:-5]
```

Cette notation fonctionne de la même façon pour les chaînes de caractères :

```python
s = 'abcdefg'
s[2:6]
```

(Bien entendu, cette notation ne peut pas fonctionner pour un dictionnaire, car les éléments des dictionnaires ne sont pas ordonnés et cela n'a donc pas de sens de dire qu'on veut les éléments du deuxième au cinquième.)

## Exercice 1 : Préfixes
Écrire et documenter une fonction `prefixes` qui prend en argument une liste d'éléments et renvoie la liste de ses préfixes, c'est-à-dire une liste de listes qui sont les début de la liste transmise en argument.

**exemple :**
```python
print(prefixes(['a', 'b', 'c', 'd', 'e']))
```
affiche
```
[[], ['a'], ['a', 'b'], ['a', 'b', 'c'], ['a', 'b', 'c', 'd'], ['a', 'b', 'c', 'd', 'e']]
 ```

<!-- grade_id: cell-3fb7b59d5305d342, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-abbb5afa991570aa, points: 2, grade: true, solution: false -->
```python
assert(prefixes(['a', 'b', 'c', 'd', 'e']) 
       == [[], ['a'], ['a', 'b'], ['a', 'b', 'c'], ['a', 'b', 'c', 'd'], ['a', 'b', 'c', 'd', 'e']])
assert(prefixes([]) == [[]])
```

## Exercice 2 : Extrait de matrice
Écrire et documenter `triangle` une fonction qui prend en argument une liste de listes, supposée carrée, et renvoie le triangle inférieur.

**exemple :**
```
print(triangle([[1,2,3,4], [5,6,7,8], [9,10,11,12], [13,14,15,16]]))
```
affiche
```
[[1], [5, 6], [9, 10, 11], [13, 14, 15, 16]]
```

<!-- grade_id: cell-8b024b717b9f6844, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-bf935c6b4fee7bd9, points: 2, grade: true, solution: false -->
```python
assert(triangle([[1,2,3,4],[5,6,7,8],[9,10,11,12],[13,14,15,16]])
      == [[1], [5, 6], [9, 10, 11], [13, 14, 15, 16]])
assert(triangle([[]]) == [[]])
```

## Exercice 3 : Extraire des colonnes pour un tableur
On dispose d'un document au format `csv` utilisant le séparateur `,` et qui possède ... colonnes. On souhaite créer un nouveau document qui ne contient que les premières colonnes.

1. Écrire et documenter une fonction `ecrire_ligne` qui prend une liste d'éléments en argument et un descripteur de fichier tel qu'obtenu par `open` et écrit dans le fichier la liste de ces éléments, séparés par des `;` et suivie d'un retour à la ligne.

<!-- grade_id: cell-4f07301780327119, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-40bae706d240958c, points: 2, grade: true, solution: false -->
```python
f = open('titi0', 'w')
ecrire_ligne([], f)
f.close()

!diff toto0 titi0 >/dev/null || echo -e '\e[07;31merreur pour le fichier titi0\e[0;m'   ## pas du python
!rm titi0 2>/dev/null ## pas du python



f = open('titi1', 'w')
ecrire_ligne(['a', 'b', 'c', 'd'], f)
f.close()

!diff toto1 titi1 >/dev/null || echo -e '\e[07;31merreur pour le fichier titi1\e[0;m'   ## pas du python
!rm titi1 2>/dev/null ## pas du python



f = open('titi2', 'w')
ecrire_ligne([1, 2, 3, 4, 5], f)
f.close()

!diff toto2 titi2 >/dev/null || echo -e '\e[07;31merreur pour le fichier titi2\e[0;m'   ## pas du python
!rm titi2 2>/dev/null ## pas du python
```

2. Écrire et documenter une fonction `colonnes` qui prend en argument le nom d'un fichier au format `csv` avec séparateur `;` et le nombre de colonnes à extraire et écrit le résultat dans le fichier `resultat.csv` (en utilisant la fonction `ecrire_ligne`).

<!-- grade_id: cell-235cf063e91f3f8e, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-894d850dd71b6c7b, points: 2, grade: true, solution: false -->
```python
colonnes('voitures.csv', 3)

!diff resultat.csv resultat3.csv || echo -e '\e[07;31merreur pour le paramètre 3\e[0;m'   ## pas du python
!rm resultat.csv 2>/dev/null ## pas du python



colonnes('voitures.csv', 0)

!diff resultat.csv resultat0.csv || echo -e '\e[07;31merreur pour le paramètre 0\e[0;m'   ## pas du python
!rm resultat.csv 2>/dev/null ## pas du python
```

3. Écrire et documenter une fonction `gruyere` qui prend en argument le nom d'un fichier au format `csv` avec séparateur `;` et deux entiers `debut` et `fin` positifs ou nuls, et écrit dans le fichier `resultat.csv` les données, avec suppression des colonnes entre `debut` compris et `fin` non compris (en utilisant la fonction `ecrire_ligne`).

<!-- grade_id: cell-5d0c5e5a94f95b45, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-f57b8dbc9b8d8d52, points: 2, grade: true, solution: false -->
```python
gruyere('voitures.csv', 3, 7)

!diff resultat.csv gruyere37.csv || echo -e '\e[07;31merreur pour les paramètres 3 et 7\e[0;m'   ## pas du python
!rm resultat.csv 2>/dev/null ## pas du python


gruyere('voitures.csv', 5, 5)

!diff resultat.csv gruyere55.csv || echo -e '\e[07;31merreur pour les paramètres 5 et 5\e[0;m'   ## pas du python
!rm resultat.csv 2>/dev/null ## pas du python
```

## Copie superficielle
On a vu que prendre une tranche revient à recopier des données. **Attention** : cette copie est superficielle:

```python
L = [ [1, 2, 3], [4, 5, 6], [7, 8, 9], [10, 11, 12]]
M = L[2:]
M[0][1] = 'pouet!!!'
L
```

Le schéma suivant illustre la situation et permet de comprendre ce qui se passe :
![](copie_superficielle.png)

Pour faire une copie profonde, il faut utiliser la fonction `deepcopy` du module `copy`. (Nous reviendrons dessus dans un autre TP.)

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* syntaxe pour une tranche : `L[i:j]`, `L[i:]`, `L[:j]`, `L[:]`
* caractère superficiel des copies par tranches
</div>