# ITC - MPSI
---

# TP19 : Complexité expérimentale
Dans ce TP, on va essayer de mesurer des complexités de manière expérimentale.

Pour cela, on a besoin de quelques outils pour:
* mesurer le temps
* représenter graphiquement des couples de valeurs

### Exercice 0a : mesurer le temps
Documentation de la fonction `time` du module `time`:

```python
import time
help(time.time)
```

Afficher le temps mis pour calculer le produit d'un entier avec lui-même, pour les entiers inférieurs à 1000000 (on ne demande pas un affichage de ces entiers, contentez-vous d'affecter leur valeur à une variable, toujours la même).

<!-- grade_id: cell-7da6576a80e38785, points: 2, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

### Exercice 0b : tracer des courbes
Si le module `matplotlib` n'est pas installé, vous devez décommenter la cellule suivante et l'exécuter, puis redémarrer le noyau.

```python
#!pip install matplotlib
```

Le module `matplotlib` contient un sous module `pyplot` que nous allons importer :

```python
import matplotlib.pyplot as plt
```

Ce sous module contient une fonction `plot` à laquelle on peut fournir deux listes de nombres de mêmes longueurs et qui renvoie un objet graphique représentant les couples de points dont l'abscisse se trouve dans la première liste et l'ordonnée dans la deuxième, dans les cases de même indice; et une fonction `show` qui permet d'afficher tous les objets graphiques créés après son appel précédent, voir par exemple :

```python
x = [1, 2, 3, 4, 5]
y = [3, 7, 2, 10, 12]
plt.plot(x, y)                # sans nom pour la courbe
plt.show()

x = [1, 2, 3, 4, 5]
y = [3, 7, 2, 10, 12]
plt.plot(x, y, label = 'a')   # avec un nom pour la courbe

x = [5, 3, 2, 1]
y = [1, 2, 3, 4]
plt.plot(x, y, label = 'b')   # avec un nom pour la courbe
plt.legend(loc='upper right') # pour placer la légende
plt.show()
```

Tracer sur un même graphique la courbe de $n\mapsto n\log_2 n$ et $n\mapsto n^2$, pour les entiers $n$ entre 1 et 100000, avec une légende. (Le module `math` contient une fonction `log2`.)

<!-- grade_id: cell-efccf27b1c7fca72, points: 2, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

### Exercice 1 : tri bulle
On rappelle le code du tri bulle donné en cours:

```python
def tri_bulle(liste) :
    """liste : liste de nombres
    sortie : la liste est triée en place"""
    permutation = True
    n = len(liste)
    while permutation :
        permutation = False
        for i in range(n-1):
            if liste[i] > liste[i+1] : # si ce n'est pas dans le bon ordre, on permute
                liste[i], liste[i+1] = liste[i+1], liste[i]
                permutation = True
        n = n - 1
```

On a montré en cours que la complexité dans le pire des cas de ce tri est en $O(n^2)$, où $n$ est la taille de l'entrée. Nous allons constaté expérimentalement que c'est également la complexité en moyenne.

La fonction `randint` du module `random` permet d'obtenir un entier aléatoire (utilisez `help` pour avoir plus d'explications). On va utiliser cette fonction pour fabriquer des entrées aléatoires pour notre algorithmes.

1. Écrire et documenter une fonction `liste_alea` qui prend en argument deux entiers `n` et `m` et renvoie une liste de `n` entiers aléatoires entre 0 et `m` compris.

<!-- grade_id: cell-eb55a0669d60b218, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-69a5652d51cd685f, points: 2, grade: true, solution: false -->
```python
assert ( len(liste_alea(50, 30)) == 50 ) 
```

2. En tirant 10 listes aléatoires par longueur, avec des entiers entre 0 et 1000, afficher la complexité expérimentale de `tri_bulle` pour les longueurs de 1 à 1000, par pas de 100 (c'est un peu long, c'est normal). Utiliser la variable `ord` pour la liste des ordonnées.

_Nota_ : la fonction `range` peut prendre un troisième paramètre qui est le pas.

<!-- grade_id: cell-d131523dad32b8e0, points: 2, grade: true, solution: true -->
```python
abs = range(1, 1000, 100)
ord = []

# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

3. La courbe ressemble (sauf erreur de votre part) à la courbe de $n\mapsto n^2$. Nous allons essayer de le prouver expérimentalement en encadrant cette courbe par des courbes de multiples de $n\mapsto n^2$.
Essayez d'encadrer au plus près votre courbe:

<!-- grade_id: cell-604467efbb28258d, points: 2, grade: true, solution: true -->
```python
plt.plot(abs, ord, label='tri bulle')

# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit

plt.legend(loc='upper right')
plt.show()
```

### Exercice 2 : quicksort
On reprend maintenant l'algorithme de quicksort. On a vu que dans le pire des cas, sa complexité est en $O(n^2)$. Qu'en est-il en pratique ?

L'algorithme :

```python
def quicksort(lst):
    tri_rapide(lst, 0, len(lst)-1)

def tri_rapide(lst, debut, fin):
    """ lst : liste dde nombres
    debut, fin : deux indices tq debut <= fin < longueur de la liste
    action : liste triee en place entre debut et fin compris"""

    if debut >= fin:
        return
    pivot = partitionner(lst, debut, fin)
    tri_rapide(lst, debut, pivot-1)
    tri_rapide(lst, pivot+1, fin)

def partitionner(lst, debut, fin):
    pivot = lst[debut] # à modifier pour que ce soit efficace
    inf = debut + 1
    sup = fin

    while True:
        while lst[sup] >= pivot and sup > debut:
            sup -= 1
        while lst[inf] < pivot and inf < fin:
            inf += 1
        if inf >= sup:
            break
        lst[inf], lst[sup] = lst[sup], lst[inf]

    lst[debut], lst[sup] = lst[sup], lst[debut]

    return sup


```

1. En tirant 30 listes aléatoires par longueur, avec des entiers entre 0 et 1000, afficher la complexité expérimentale de `quicksort` pour les longueurs de 1 à 4000, par pas de 200 (c'est un peu long, c'est normal). Utiliser la variable `ord2` pour la liste des ordonnées.

<!-- grade_id: cell-1a5154f9d7a5d63e, points: 2, grade: true, solution: true -->
```python
abs2 = range(1, 4000, 200)
ord2 = []

# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

2. Tracer maintenant sur un même graphique les complexités expérimentale du tri bulle et du tri rapide.

<!-- grade_id: cell-c0abfe48a46fa58b, points: 2, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

### Exercice 3 : mélange de tris
En fait pour les petites listes, le tri bulle est souvent plus rapide que quicksort, car les petites listes sont presque triées.

1. Écrivez une fonction `melange_tris` qui prend en argument une liste et trie la liste en place en agissant comme le tri rapide si la liste est de longueur strictement supérieur à 5, et comme le tri bulle sinon.

<!-- grade_id: cell-bf42b2367c868b6f, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-b47099c737b0052a, points: 2, grade: true, solution: false -->
```python
import random
random.seed(5)
lst = [random.randint(1,100) for i in range(100)]
copie = lst[:]
melange_tris(lst)
assert (lst == sorted(copie))
```

2. En tirant 30 listes aléatoires par longueur, avec des entiers entre 0 et 1000, afficher la complexité expérimentale de `melange_tris` pour les longueurs de 1 à 4000, par pas de 200 (c'est un peu long, c'est normal). Utiliser la variable `ord3` pour la liste des ordonnées.

<!-- grade_id: cell-550cbdb7829671ac, points: 2, grade: true, solution: true -->
```python
abs3 = range(1, 4000, 200)
ord3 = []

# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

3. Tracer maintenant sur un même graphique les complexités expérimentale du tri bulle, du tri rapide et du mélange des deux.

<!-- grade_id: cell-61795e19d3349344, points: 2, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

### Exercice 4 : ajout dans une liste `python`
Le but de cet exercice est de constater que l'ajout d'un élément dans une liste python par la méthode `append` n'est pas toujours opération qui peut être considérée en temps constant.

Pour celà, vous allez suivre les étapes suivantes:
* créer une liste de 10 listes vides,
* ajouter un à un les 500000 premiers entiers à chacune des listes précédente, en utilisant la méthode `append`, et en chronométrant le temps mis pour ajouter l'entier $i$ à l'ensemble des listes,
* dessiner la courbe qui prend en abscisse les 500000 premiers entiers et en ordonnée le temps mis pour ajouter chacun de ces entiers aux dix listes.

<!-- grade_id: cell-bd1853ba1045d5af, points: 2, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```