# ITC - MPSI
---

# TP2 : conditionnelles et boucles conditionnelles

Rappel: ce TP est à rendre à la fin de la séance. Un autre énoncé sera mis en ligne au plus tard vendredi soir, à rendre pour le mardi soir.

## Exo 1 : nombre de chiffres

En `python`, le quotient de la division euclidienne entre entiers s'obtient grâce à l'opérateur `//`. Comparer les résultats des cellules suivantes.

(Un _flottant_ est le type de données qui permet de représenter les réels. Rappel : les calculs sur les flottants entraînent des erreurs d'arrondi, on ne teste donc **jamais** l'égalité (ou la non égalité) entre deux flottants.)

```python
1 / 2   ## est un flottant
```

```python
1 // 2  ## est un entier
```

```python
4 / 2   ## est un flottant
```

```python
4 // 2  ## est un entier
```

À partir d'un entier `n`, on peut donc construire l'entier obtenu en effaçant le dernier chiffre de `n` comme le quotient entier de `n` par 10.

Documentez et écrivez une fonction `nb_de_chiffres` qui prend en argument un entier supposé positif (sans avoir à le vérifier) et renvoie son nombre de chiffres (par convention on considèrera que le nombre de chiffres de l'entier 0 est 0).

<!-- grade_id: cell-70f5365b378885ec, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-17dd89ab61455287, points: 1, grade: true, solution: false -->
```python
assert(nb_de_chiffres(0) == 0)
assert(nb_de_chiffres(1) == 1)
assert(nb_de_chiffres(9) == 1)
assert(nb_de_chiffres(10) == 2)
assert(nb_de_chiffres(99) == 2)
assert(nb_de_chiffres(100) == 3)
assert(nb_de_chiffres(865462) == 6)
```

## Exo 2 : nombre miroir

Un nombre est le miroir d'un autre si ses chiffres sont dans l'ordre inverse. Exemple : 5 est son propre miroir, 721 est le miroir de 127. Les miroirs des multiples de 10 apparaîtront bien entendu sans les 0 initiaux, ainsi le miroir de 120 est 21.

Notez que si on veut reconstituer un nombre à trois chiffres dont on connaît le chiffre des centaines $c$, le chiffre des dizaines $d$ et le chiffre des unités $u$, il suffit d'exécuter les étapes suivantes pour récupérer le nombre :
$$c\xrightarrow{\times 10} 10c\xrightarrow{+d}10c+d\xrightarrow{\times 10}100c+10d\xrightarrow{+u}100c+10d+u$$

Le chiffre des unités d'un entier est le reste de la division euclidienne de cet entier par 10 et s'obtient en `python`grâce à l'opérateur `%`: ainsi `37 % 10` s'évalue à `7`.

a. Documentez et écrivez une fonction `miroir` qui prend en argument un entier supposé positif ou nul (sans avoir à le vérifier) et renvoie son miroir.

<!-- grade_id: cell-4d14ca67f6a12f2e, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-741babbf4127c2c9, points: 1, grade: true, solution: false -->
```python
assert(miroir(0) == 0)
assert(miroir(1) == 1)
assert(miroir(10) == 1)
assert(miroir(127) == 721)
assert(miroir(1200) == 21)
assert(miroir(943742) == 247349)
```

b. Un nombre est un _palindrome_ si on peut le lire dans les deux sens (et qu'on obtient le même nombre). Exemple : 5, 12721 et 127721 sont des palindromes, 12723 et 1210 n'en sont pas.

Documentez et écrivez une fonction `est_palindrome` qui prend en argument un entier supposé positif ou nul (sans avoir à le vérifier) et teste si c'est un palindrome (c'est-à-dire renvoie un booléen: `True` si l'entier est un palindrome et `False` sinon). N'oubliez pas de traiter le cas où `n` est divisible par `0` mais non nul.

<!-- grade_id: cell-37a37f463f0dab95, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-e7685d93c6f63bf3, points: 1, grade: true, solution: false -->
```python
assert(est_palindrome(0))
assert(est_palindrome(5))
assert(est_palindrome(12721))
assert(est_palindrome(127721))
assert(not est_palindrome(12723))
assert(not est_palindrome(765732))
assert(not est_palindrome(1210))
```

## Exo3 : plus petit palindrome immédiatement supérieur

On veut maintenant trouver, pour un entier `n` donné, le plus petit entier supérieur ou égal à `n` qui est un palindrome.

Comme vous pouvez le constater dans la cellule de test précédente, on peut prendre la négation d'un booléen grâce à l'opérateur `not` placé devant. Exécutez les cellules suivantes pour mieux comprendre:

```python
not True
```

```python
not False
```

```python
not 1 < 5
```

```python
est_palindrome(127)
```

```python
not est_palindrome(127)
```

Documentez et écrivez la fonction `approximation_palindromique` (ce n'est pas un terme officiel) qui à un entier `n` associe le plus petit entier supérieur ou égal à `n` qui est un palindrome.

<!-- grade_id: cell-3937aa4084e22c5c, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-8e0597de0bfa1606, points: 0, grade: true, solution: false -->
```python
assert(approximation_palindromique(5) == 5)
assert(approximation_palindromique(10) == 11)
assert(approximation_palindromique(127) == 131)
assert(approximation_palindromique(453624) == 454454)
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>

* La division euclidienne entre deux entiers s'obtient par les opérateurs `//` (quotient) et `%` (reste).
* L'opérateur `/` fait une division entre flottants (même si on lui fournit des entiers).
* En `python`, les constantes booléennes sont `True` et `False`. La négation d'un booléen s'exprime avec l'opérateur `not`.
* Un _test_ est une fonction qui renvoie un booléen.
</div>