# ITC - MPSI
---

# TP20 : Matrices d'adjacence
Le but de ce TP est de manipuler des graphes sous forme de matrices d'adjacence.

Une _matrice_ sera donnée sous forme de liste de listes. 

### Exercice 1: comptages divers
1. Écrire et docmenter une fonction `degre_sortant` qui prend en argument la matrice d'adjacence d'un graphe et le numéro d'un sommet et renvoie le degré sortant de ce sommet.

<!-- grade_id: cell-577393c33a9791cb, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-c4a53afeb017c209, points: 2, grade: true, solution: false -->
```python
M = [ 
    [1, 0, 1, 0],
    [1, 1, 0, 0],
    [0, 0, 0, 0],
    [0, 1, 0, 0]
]
assert (degre_sortant(M, 0) == 2)
assert (degre_sortant(M, 1) == 2)
assert (degre_sortant(M, 2) == 0)
assert (degre_sortant(M, 3) == 1)
```

2. Écrire et docmenter une fonction `degre_entrant` qui prend en argument la matrice d'adjacence d'un graphe et le numéro d'un sommet et renvoie le degré entrant de ce sommet.

<!-- grade_id: cell-c4eb08eeda3abcbc, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-c4494e7a196ebbb3, points: 2, grade: true, solution: false -->
```python
M = [ 
    [1, 0, 1, 0],
    [1, 1, 0, 0],
    [0, 1, 0, 0],
    [0, 1, 0, 0]
]
assert (degre_entrant(M, 0) == 2)
assert (degre_entrant(M, 1) == 3)
assert (degre_entrant(M, 2) == 1)
assert (degre_entrant(M, 3) == 0)
```

3. Écrire et documenter une fonction `boucles` qui prend en argument la matrice d'adjacence d'un graphe et renvoie le nombre de boucles de ce graphe.

<!-- grade_id: cell-15f22478368a4e6b, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-075adc6de00bb29e, points: 2, grade: true, solution: false -->
```python
M = [ 
    [1, 0, 1, 0],
    [1, 1, 0, 0],
    [0, 1, 0, 0],
    [0, 1, 0, 1]
]
assert (boucles(M) == 3)

t = 10
M = [ [0]*t for _ in range(t) ]
for i in range(t):
    M[i][i] = 1
assert (boucles(M) == t)

M = [ [1]*t for _ in range(t) ]
for i in range(t):
    M[i][i] = 0
assert (boucles(M) == 0)
```

4. Écrire et documenter une fonction `arcs` qui prend en argument la matrice d'adjacence d'un graphe et renvoie le nombre d'arcs de ce graphe.

<!-- grade_id: cell-583fbead88cec0ae, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-39bfbf2a1200a1a6, points: 2, grade: true, solution: false -->
```python
M = [ 
    [1, 0, 1, 0],
    [1, 1, 0, 0],
    [0, 1, 0, 0],
    [0, 1, 0, 1]
]
assert (arcs(M) == 7)

t = 10
M = [ [0]*t for _ in range(t) ]
assert (arcs(M) == 0)

M = [ [1]*t for _ in range(t) ]
assert (arcs(M) == t*t)
```

### Exercice 2 : distance
On définit la _distance_ d'un sommet $i$ vers un sommet $j$ comme la longueur du plus court chemin allant de $i$ à $j$. On la note $d(i,j)$.

1. Si un graphe possède $n$ sommets, quelle est a la longueur maximale d'un chemin qui ne passe pas deux fois par le même sommet (sauf éventuellement ses extrêmités qui peuvent être égales) ?

<!-- grade_id: cell-1ae6c57e486b6503, points: 1, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

S'il existe un chemin de $i$ vers $j$ dans un graphe $G=(S,A)$, on a la relation de récurrence suivante:

$$d(i,j) = 1 + \min\limits_{k\in S\mid (i,k)\in A} d(k,j).$$

Il n'est pourtant pas aisé d'écrire une fonction donnant la distance entre deux sommets, car il n'est pas évident de poser un cas initial, et on tombe vite dans un cycle qui entraîne un débordement de pile:
* pour savoir si on peut aller de 1 à 3, on vérifie si on peut aller de 1 à 2 et de 2 à 3,
* pour savoir si on peut aller de 2 à 3, on vérifie si on peut aller de 2 à 1 et de 1 à 3,
* pour savoir si on peut aller de 1 à 3, on vérifie si on peut aller de 1 à 2 et de 2 à 3,
* pour savoir si on peut aller de 2 à 3, on vérifie si on peut aller de 2 à 1 et de 1 à 3,
* pour savoir si on peut aller de 1 à 3, on vérifie si on peut aller de 1 à 2 et de 2 à 3,
* etc.

Pour s'en sortir, on va ajouter à notre fonction un marqueur qui sera en quelque sorte le témoin du nombre de sommets qu'on a déjà visités. Si ce marqueur est supérieur au nombre de sommets du graphe, c'est qu'on s'est embarqué dans un cycle, et on arrête le parcours.

Ainsi on peut définir
$$d'(i, j, \mu) = \begin{cases} d(i, j)\text{ si }\mu <= |S|\\
|S|\text{ sinon}
\end{cases}$$

Et on a la relation de récurrence:
$$d'(i,j,\mu) = 1 + \min\limits_{k\in S\mid (i,k)\in A} d'(k,j,\mu+1).$$


2. Écrire et documenter une fonction `distance` qui prend en argument la matrice d'adjacence d'un graphe, deux sommets de ce graphe (donnés sous la forme de leurs indices dans cette matrice) **tels qu'il existe au moins un chemin du premier vers le deuxième** et un marqueur jouant le rôle qu'on vient de décrire et renvoie la distance du premier sommet vers le deuxième.

<!-- grade_id: cell-f4deb0bf9f6fc510, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-2f22d4a60f5c0dc5, points: 2, grade: true, solution: false -->
```python
M = [ 
    [1, 0, 1, 0],
    [1, 1, 0, 0],
    [0, 1, 0, 0],
    [0, 1, 0, 1]
]
assert (distance(M, 1, 1, 0) == 0)
assert (distance(M, 0, 2, 0) == 1)
assert (distance(M, 0, 1, 0) == 2)
assert (distance(M, 3, 2, 0) == 3)
```

3. Quelle sera la valeur renvoyée par votre fonction s'il n'existe pas de chemin entre les deux sommets fournis en entrée?

<!-- grade_id: cell-db7ea525e414e068, points: 1, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

### Exercice 3 : sommets atteignables
En s'insispirant de la mise en œuvre de la fonction `distance` précédente, écrire et documenter une fonction `atteignables` qui prend en argument la matrice d'adjacence d'un graphe, un sommet de ce graphe, et une liste de sommets et renvoie la liste des sommets atteignables à partir du sommet fourni.

<!-- grade_id: cell-76c73ad08488d037, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-9573495c0fbd52db, points: 2, grade: true, solution: false -->
```python
M = [ 
    [1, 0, 1, 0, 0],
    [1, 1, 0, 0, 0],
    [0, 1, 0, 0, 0],
    [0, 1, 0, 1, 0],
    [0, 0, 0, 0, 0]
]

assert (atteignables(M, 4, []) == [4])
assert (sorted(atteignables(M, 0, [])) == [0, 1, 2])
assert (sorted(atteignables(M, 3, [])) == [0, 1, 2, 3])
```

### Exercice 4 : retournement de situation
Écrire et documenter une fonction `miroir` qui prend en argument la matrice d'adjacence d'un graphe et renvoie la matrice d'adjacence du graphe obtenu en retournant les arcs.

<!-- grade_id: cell-c1475746a13f0aeb, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-68ff9a8944ef1ed2, points: 2, grade: true, solution: false -->
```python
M = miroir ([ 
    [1, 0, 1, 0],
    [1, 1, 0, 0],
    [0, 1, 0, 0],
    [0, 1, 0, 1]
])
assert( M == [[1, 1, 0, 0], [0, 1, 1, 1], [1, 0, 0, 0], [0, 0, 0, 1]])
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* On peut représenter un graphe orienté par une matrice appelée _matrice d'adjacence_ du graphe.
</div>