# ITC - MPSI
---

# TP21 : Listes d'adjacence

### Exercice : des matrices aux listes
Écrire et documenter une fonction `mat_to_list` qui prend en argument une matrice d'adjacence d'un graphe orienté et renvoie une liste contenant les listes d'adjacence représentant le même graphe (avec les numéros de sommets correspondants).

<!-- grade_id: cell-5ee2537ceefd1221, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-184f2479fb748552, points: 2, grade: true, solution: false -->
```python
assert ( mat_to_list([[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 1, 0, 1, 0, 0, 0, 0, 0, 0],
 [0, 0, 1, 1, 0, 0, 0, 0, 1, 0, 0],
 [0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0],
 [0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0],
 [0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0],
 [0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0],
 [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0]]) == [[], [2, 4], [2, 3, 8], [1], [5], [2], [7], [6], [9], [], [3]] )
```

### Exercice : des listes aux matrices
Écrire et documenter une fonction `list_to_mat` qui prend en argument une liste contenant les listes d'adjacence représentant un graphe orienté et renvoie la matrice d'adjacence du même graphe (avec les numéros de sommets correspondants).

<!-- grade_id: cell-149152dc7a23c5f6, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-627b791871f21dd6, points: 2, grade: true, solution: false -->
```python
assert ( list_to_mat([ [], [2, 4], [2, 3, 8], [1], [5], [2], [7], [6], [9], [], [3]])
        == [[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 1, 0, 1, 0, 0, 0, 0, 0, 0],
 [0, 0, 1, 1, 0, 0, 0, 0, 1, 0, 0],
 [0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0],
 [0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0],
 [0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0],
 [0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0],
 [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0]] )
```

### Exercice : Représentation graphique
On peut utiliser la classe `Digraph` du module `graphviz` pour obtenir un dessin d'un graphe orienté.

Commencez par installer le paquet graphviz en exécutant la cellule suivante. Et pensez à redémarrez le noyau de votre notebook.

```python
!pip install --user graphviz
```

On peut alors dessiner un graphe orienté de la façon suivante:

```python
from graphviz import Digraph

graphe = Digraph()   ## construction de l'objet
# ajout des noeuds avec des étiquettes
graphe.node('1')
graphe.node('2')
graphe.node('toto')

# ajout d'arcs
graphe.edges([('1','2'), ('toto', '1'), ('2', '2')])

# affichage de l'objet
graphe
```

Écrire et documenter une fonction `dessin_graphe` qui prend en argument un graphe donné sous forme de listes d'adjacence et renvoie l'objet qui représente graphiquement le graphe. Le sommet numéroté $i$ devra porter l'étiquette `si`.

<!-- grade_id: cell-2d058f7dee6ca9c0, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-01b76c79b3e8f37c, points: 2, grade: true, solution: false -->
```python
dessin_graphe([ [], [2, 4], [2, 3, 8], [1], [5], [2], [7], [6], [9], [], [3]])
```

### Exercice 4 : Retournement de situation
Écrire et documenter une fonction `miroir` qui prend en argument les listes d'adjacence d'un graphe orienté et renvoie les listes d'adjacence du graphe obtenu en retournant les arcs.

À vous de faire des tests avec des dessins !

Describe the task here!