# ITC - MPSI
---

# TP22 : Exploration d'un vrai graphe
Le but de ce TP est d'explorer un graphe provenant de données réelles. Il s'agit ici d'explorer les followers à partir du compte [@ENSdeLyon](https://twitter.com/ENSdeLyon). (J'ai emprunté le fichier à un collègue.)

Les données se trouvent dans le fichier `ENSdeLyon.graph`.

Elles ont le format suivant:
* la première ligne contient le nombre de comptes répertoriés, soit $n$;
* les $n$ lignes suivantes contiennent les noms des comptes;
* la ligne suivante contient le nombre de liens de type `followers`, soit $f$;
* les $f$ lignes suivantes contiennent chacune deux entiers : l'indice du compte qui suit et l'indice du compte qui est suivi (les comptes sont numérotés à partir de 0).

### Exercice 1 : Taille
En allant chercher les informations dans le fichier (avec du code), combien y a-t-il de comptes répertoriés dans le fichier `ENSdeLyon.graph`(n'oubliez pas le chemin d'accès complet si vous êtes sur le serveur de TP) ? Combien y a-t-il de liens ?   

Si on considère le graphe orienté dont les sommets sont les comptes et les arcs les relations suiveur/suivi, quelle est la proportion d'arcs dans le graphe par rapport au plus grand nombre d'arcs possibles?

(Laissez le code que vous écrivez pour avoir la réponse dans la cellule suivante, en plus de compléter les valeurs.)

<!-- grade_id: cell-dc5d90f19e5a0fee, points: 0, grade: false, solution: true -->
```python
nb_accounts = 0    #à modifier
nb_follows  = 0    #à modifier
proportion  = 0    #à modifier
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-f0e94fa47f0a4c20, points: 1, grade: true, solution: false -->
```python

```

### Exercice 2 : Coder le graphe
En principe, la réponse à l'exercice précédent a dû vous convaincre que ce n'est pas une bonne idée de représenter ce graphe avec une matrice d'adjacence. On va donc représenter le graphe avec une liste d'adjacence.

Écrire et documenter une fonction `graphe_twitter` qui prend en argument un nom de fichier au même format que `ENSdeLyon.graph` et renvoie un couple dont le premier argument est la liste des comptes dans leur ordre d'apparition dans le fichier, et le deuxième argument la liste d'adjacence correspondante.

<!-- grade_id: cell-26d2a14ba5a0d9a4, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-87abb7313c9bb8ce, points: 2, grade: true, solution: false -->
```python
sommets, adj = graphe_twitter('ENSdeLyon.graph')
assert(sommets[17] == 'MarioGhossoub')
assert(adj[17] == [5443,1113,5720,8305,821,8370,5658,1327,3680,5875,7258,3783,2956,4812,3541,1214])
assert(sommets[2022] == 'cosyeli')
assert(adj[2022] == [])
```

### Exercice 3 : Sous-graphe induit
Si $G=(S,A)$ est un graphe et $S'\subseteq S$ un ensemble de sommets de $G$, alors le sous-graphe de $G$ induit par $S'$ est le graphe $G'=(S',A')$ où $$A'=\{(s,t)\in A\mid s\in S'\text{ et }t\in S'\}\:.$$

Écrire et documenter une fonction `sous_graphe_induit` qui prend en argument une liste d'adjacence telle que récupérée par la fonction `graphe_twitter` et un entier $n$ et extrait le sous-graphe induit par les sommets $\{0,1,\dots,n-1\}$.

<!-- grade_id: cell-06d80a465d8b4749, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-3cc79141b7a6a74f, points: 2, grade: true, solution: false -->
```python
sommets, adj = graphe_twitter('ENSdeLyon.graph')
adj2 = sous_graphe_induit(adj, 2500)
assert(adj2[17] == [1113, 821, 1327, 1214])
assert(adj2[2022] == [])
```

### Exercice 4 : Degrés sortants
1. Écrire et documenter une fonction `degre_sortant_moyen` qui prend en argument une liste d'adjacences comme récupérée de la fonction `graphe_twitter`, et renvoie le degré sortant moyen d'un sommet dans le graphe correspondant.

<!-- grade_id: cell-d75e904380f37757, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

```python
sommets, adj = graphe_twitter('ENSdeLyon.graph')
degre_sortant_moyen(adj)
```

2. Écrire et documenter une fonction `degre_sortant_max` qui prend en argument une liste de sommets et une liste d'adjacences comme récupérées de la fonction `graphe_twitter`, et renvoie le nom du sommet ayant le degré sortant maximal (si ce sommet n'est pas unique, le nom du premier ayant cette propriété dans la liste des sommets).

<!-- grade_id: cell-5b15e574e25f0dc1, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-cd34df32029c4448, points: 2, grade: true, solution: false -->
```python
sommets, adj = graphe_twitter('ENSdeLyon.graph')
assert(degre_sortant_max(sommets, adj) == 'EditionsScpo')
sommets800 = sommets[:800]
adj800 = sous_graphe_induit(adj, 800)
assert(degre_sortant_max(sommets800, adj800) == 'CbdLes')
```

### Exercice 5 : Degré entrant moyen
Écrire et documenter une fonction `degre_entrant_moyen` qui prend en argument une liste d'adjacences comme récupérée de la fonction `graphe_twitter`, et renvoie le degré entrant moyen d'un sommet dans le graphe correspondant. Faites en sorte que l'exécution de votre fonction ne soit pas trop longue.

<!-- grade_id: cell-4d9233b2ca2d4148, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

Tester le temps d'exécution sur un sous-graphe de taille 200, puis de taille 1000 avant de lancer les tests sur des sous-graphes plus grands (chez moi c'est respectivement moins de 100 micro-secondes et moins de 300 micro-secondes, sur le serveur de TP):

```python
sommets, adj = graphe_twitter('ENSdeLyon.graph')

sommets200 = sommets[:200]
adj200 = sous_graphe_induit(adj, 200)

sommets1000 = sommets[:1000]
adj1000 = sous_graphe_induit(adj, 1000)
```

```python
%%time
from math import fabs

assert(fabs(degre_entrant_moyen(adj200) - 0.23) < 0.00001)
```

```python
%%time
from math import fabs

assert(fabs(degre_entrant_moyen(adj1000) - 1.275) < 0.00001)
```

<!-- grade_id: cell-084c11eed0af4ed7, points: 2, grade: true, solution: false -->
```python
sommets, adj = graphe_twitter('ENSdeLyon.graph')
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* définition du sous-graphe induit par un ensemble de sommets
* la somme des degrés entrants est égale à la somme des degrés sortants
</div>