# ITC - MPSI
---

# TP3 : manipulation de listes

Dans ce TP, nous allons étudier quelques algorithmes élémentaires concernant les listes d'éléments. Ce sera également l'occasion de rappeler quelques instructions de `python` et de voir ou revoir les listes.

## Listes en `python`

En `python`, une _liste_ est une suite finie et ordonnée d'éléments. Ces éléments n'ont pas besoin d'être du même type :
```python
liste_entiers = [1, 4, 8, -2]
pi = 3.14
liste_quelconque = [-17, 'toto', 3.4, pi, True]
```

En fait une liste peut même posséder un élément qui est lui-même une liste :
```python
liste_compliquee = ['titi', ['a', 3, 1.2], 9.5]
```
mais pour ce TP nous éviterons de les utiliser.

En `python`, les éléments d'une liste sont repérés par leurs _indices_ dans la liste, et ces indices commencent **toujours** par 0. On utilise des crochets pour accéder à un élément à partir de son indice. Par exemple :

```python
L = ['a', 'b', 'c', 'd']
print(L[1])
```

```python
L[0] = 'z'
print(L)
```

On peut connaître la longueur d'une liste grâce à la fonction `len`:

```python
len(L)
```

### Exo 1 : accéder aux éléments

1. Écrire et documenter une fonction `avant_dernier` qui prend en argument une liste supposée contenir au moins deux éléments (sans avoir à le vérifier) et renvoie son avant-dernier élément.

<!-- grade_id: cell-fe0a24d825375171, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-1d0243295cabe404, points: 1, grade: true, solution: false -->
```python
assert(avant_dernier([3, 4]) == 3)
assert(avant_dernier([9, -3, 5, 17]) == 5)
assert(avant_dernier([-1, -1, -1, -1, -1]) == -1)
```

2. Écrire et documenter une fonction `ddebut` qui prend en argument une liste supposée de longueur au moins 2 et supposée contenir des entiers (sans avoir à le vérifier) et teste si ses deux premiers éléments sont égaux.

<!-- grade_id: cell-964f522edd16cb33, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-a713098df7f70a44, points: 1, grade: true, solution: false -->
```python
assert(ddebut([4,4,2,-4]))
assert(not ddebut([5,2,5]))
```

## Boucles inconditionnelles

On a vu dans le TP précédent l'instruction `while` qui permet d'exécuter un bloc d'instructions en boucle tant qu'une certaine condition est vraie. Il existe en `python` une instruction qui permet d'exécuter un bloc d'instructions pour chaque élément d'une liste : `for`.

```python
# liste est une liste
for x in liste:
    #code
```
Le bloc est alors réexécuté pour chaque élément de la liste.

Exemple :

```python
L = [34, -15, 6, 27]
for x in L :
    if x%2 == 0 :
        print(str(x) + " pair")  # str(i) permet d'obtenir une chaînes de caractères contenant x
    else:
        print(str(x) + " impair")
```

### Exo 2 : parcourir les éléments d'une liste

1. Écrire et documenter une fonction `somme` qui prend en argument une liste de nombres éventuellement vide (sans avoir à le vérifier) et renvoie sa somme (sans utiliser la fonction `sum` de `python`).

<!-- grade_id: cell-3f0bc28428fd40dc, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-a09100ba64c5218e, points: 1, grade: true, solution: false -->
```python
import random

assert(somme([]) == 0)

liste1 = random.sample(range(10, 30), random.randint(5,15)) # liste aléatoire
assert(somme(liste1) == sum(liste1))
```

2. Écrire et documenter une fonction `que_des_pairs` qui prend en argument une liste d'entiers éventuellement vide (sans avoir à le vérifier) et teste si elle ne contient que des entiers pairs (rappel : le reste de la dixision euclidienne de `x` par 2 s'écrit `x%2` en `python`).

<!-- grade_id: cell-11c98315c5f4aa46, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-0379d195f63b4732, points: 1, grade: true, solution: false -->
```python
assert(que_des_pairs([]))
assert(que_des_pairs([ 2*i-20 for i in random.sample(range(10, 30), random.randint(5,15))]))
assert(not que_des_pairs([34, 22, 17, -12]))
```

3. Écrire et documenter une fonction `max_liste` qui prend en argument une liste de nombres non vide (sans avoir à le vérifier) et renvoie son maximum (sans utiliser la fonction `max` de `python`).

<!-- grade_id: cell-6aaa6daf2938418a, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-cb5bfc8c177c9828, points: 1, grade: true, solution: false -->
```python
import random

liste1 = random.sample(range(10, 30), random.randint(5,15))
assert(max_liste(liste1) == max(liste1))

liste2 = [ -i for i in random.sample(range(10, 30), random.randint(5,15)) ]
assert(max_liste(liste2) == max(liste2))
```

4. Écrire et documenter une fonction `tous_egaux` qui prend en argument une liste d'entiers non vide (sans avoir à le vérifier) et teste si tous les éléments de la liste sont égaux.

<!-- grade_id: cell-534e947566c63765, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-13e6c6ac5bc66850, points: 1, grade: true, solution: false -->
```python
import random

liste1 = random.sample(range(10, 30), random.randint(5,15))
assert(not tous_egaux(liste1))

assert(not tous_egaux([1,1,1,1,4]))
assert(not tous_egaux([4,1,1,1,1]))
assert(not tous_egaux([1,4,1,1,1]))
assert(tous_egaux([1,1,1,1]))
assert(tous_egaux([1]))
```

## Boucler sur des entiers

On peut avoir besoin de parcourir le début des entiers. On peut bien sûr écrire

```python
for i in [0, 1, 2, 3, 4, 5] :
    # code
```
mais si on veut parcourir les 500 premiers entiers, ce serait fastidieux et si on veut parcourir les `n` premiers entiers où `n` est un entier dont on ne connaît pas la valeur au moment où on écrit le programme, c'est carrément impossible.

Pour cela, on utilise une fonction `range` qui renvoie un _itérateur_ (notion que vous n'avez pas besoin de comprendre pour utiliser cette fonction), et qui permet de faire une boucle sur des entiers de la manière suivante :

```python
for i in range(50):
    print(i)
```

Notez bien que dans la boucle précédente, la variable `i` prend successivement toutes les valeurs de 0 à 49.

### Exo 3 : parcourir les indices d'une liste

1. Écrire et documenter une fonction `indice_max` qui prend en argument une liste de nombres supposée non vide (sans avoir à le vérifier) et renvoie l'indice de la première occurence du maximum de la liste. 

<!-- grade_id: cell-677c1e889db4ef73, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-9959d8e6732af605, points: 1, grade: true, solution: false -->
```python
assert(indice_max([0]) == 0)
assert(indice_max([-5, -7, -5, -10]) == 0)
assert(indice_max([4, 10, 3, 9, 10, 5]) == 1)
assert(indice_max(list(range(39))) == 38)
```

2. Écrire et documenter une fonction `indice` qui prend en argument une liste éventuellement vide d'entiers (sans avoir à le vérifier) et une valeur et renvoie l'indice de première occurence de cette valeur si elle appartient à la liste, -1 sinon. (Sans utiliser la méthode `index` de la classe `list` en `python`.)

<!-- grade_id: cell-743c35c6d7414725, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-34dc38b6843b20bb, points: 1, grade: true, solution: false -->
```python
import random

liste1 = random.sample(range(10, 30), random.randint(5,15))
n = random.randint(0, len(liste1)-1)
assert(indice(liste1, liste1[n]) == n)
assert(indice(liste1, 35) == -1)

assert(indice(list(range(10)), 5) == 5)
```

2. Écrire une fonction `croissante` qui prend en argument une liste de nombres éventuellement vide (sans avoir à le vérifier) et teste si les nombres apparaissent dans l'ordre croissant au sens large. (Il suffit de regarder les couples de nombres de gauche à droite pour le savoir.)

<!-- grade_id: cell-f0b20e3120ddd9aa, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-bb52c58d4a7601da, points: 1, grade: true, solution: false -->
```python
assert(croissante([]))
assert(croissante([2]))
assert(croissante([1,2,3,4]))
assert(croissante([-5,-4,-3,-2,-1,0]))
assert(not croissante([2,1]))
assert(not croissante([1,2,3,4,5,0]))
```

3.[difficile] Écrire et documenter une fonction `max2` qui prend en argument une liste d'entiers supposée non vide (sans avoir à le vérifier) et renvoie son deuxième maximum s'il existe, ou son maximum sinon.

<!-- grade_id: cell-fa84f52f10e698f0, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

```python
assert(max2([1,1,1,1]) == 1)
assert(max2([1,2,1,1]) == 1)
assert(max2([2,2,2,2,2,1]) == 1)
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* Une liste `python` s'écrit entre crochets `[ ]`.
* Les indices d'une liste commencent à 0.
* L'élément de la liste `L` situé dans la case `i` s'écrit `L[i]`.
* La longueur de la liste `L` s'obtient par `len(L)`.
* `for` permet de boucler sur les éléments d'une liste.
* Pour boucler sur les entiers de `0` à `n-1`, on écrit `for i in range(n):`
    
</div>

```python

```