# ITC - MPSI
---

# TP4 : Retour sur les paramètres des fonctions et boucles imbriquées
Dans ce TP nous regardons de plus près ce que deviennent les paramètres dans les fonctions, et nous faisons des boucles dans les boucles.

Quand on transmet un paramètre _numérique_ (entier ou nombre à virgule flottante) à une fonction, celui-ci est recopié et la fonction travaille sur une copie de ce paramètre. Vous pouvez regarder l'exécution du code suivant :
```python
"""fonction qui prend en argument un entier n et renvoie le triple de son
successeur, soit 3*(n+1) (fait de façon un peu tordue pour illustrer mon
propos"""
def triple_successeur(n) :
    n = n+1
    n = 3*n
    return n
    
n = 5
m = triple_successeur(n)
```
sur [python tutor](https://pythontutor.com/visualize.html#code=%22%22%22fonction%20qui%20prend%20en%20argument%20un%20entier%20n%20et%20renvoie%20le%20triple%20de%20son%0Asuccesseur,%20soit%203*%28n%2B1%29%20%28fait%20de%20fa%C3%A7on%20un%20peu%20tordue%20pour%20illustrer%20mon%0Apropos%22%22%22%0Adef%20triple_successeur%28n%29%20%3A%0A%20%20%20%20n%20%3D%20n%2B1%0A%20%20%20%20n%20%3D%203*n%0A%20%20%20%20return%20n%0A%20%20%20%20%0An%20%3D%205%0Am%20%3D%20triple_successeur%28n%29&cumulative=false&curInstr=0&heapPrimitives=nevernest&mode=display&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false) : la variable `n` du contexte global (celle qui est initialisée à l'avant-dernière ligne du code, puis fournie en argument à `triple_successeur`) fournit uniquement sa _valeur_ au paramètre de `triple_successeur`. Ce paramètre vit ensuite sa propre vie et peut désigner une autre valeur, cela n'a pas d'impact sur la variable `n` du contexte global.

Quand on transmet une liste comme paramètre à une fonction, c'est l'adresse de cette liste qui est recopiée (c'est-à-dire l'endroit où elle se trouve dans la mémoire) car la valeur d'une liste est l'adresse de sa première case (ou cellule). La fonction travaille donc sur une copie de cette adresse, et elle peut aller modifier le contenu des cases dans la mémoire. Cette modification est donc visible hors de la fonction. Observez la situation avec ce programme :
```python
"""fonction qui prend en argument une liste et recopie son premier
élément dans toute ses cases"""
def uniformisation(liste):
    for i in range(1, len(liste)) : #boucle de 1 à len(liste)-1
        liste[i] = liste[0]
    liste = [3, 2] # ligne inutile, à portée uniquement pédagogique
    
L = [1, 2, 3, 4, 5]
uniformisation(L)
```
sur [python tutor](https://pythontutor.com/visualize.html#code=%22%22%22fonction%20qui%20prend%20en%20argument%20une%20liste%20et%20recopie%20son%20premier%0A%C3%A9l%C3%A9ment%20dans%20toute%20ses%20cases%22%22%22%0Adef%20uniformisation%28liste%29%3A%0A%20%20%20%20for%20i%20in%20range%281,%20len%28liste%29%29%20%3A%20%23boucle%20de%201%20%C3%A0%20len%28liste%29-1%0A%20%20%20%20%20%20%20%20liste%5Bi%5D%20%3D%20liste%5B0%5D%0A%20%20%20%20liste%20%3D%20%5B3,%202%5D%20%23%20ligne%20inutile,%20%C3%A0%20port%C3%A9e%20uniquement%20p%C3%A9dagogique%0A%20%20%20%20%0AL%20%3D%20%5B1,%202,%203,%204,%205%5D%0Auniformisation%28L%29&cumulative=false&curInstr=0&heapPrimitives=nevernest&mode=display&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false) : la variable `L` du contexte global et le paramètre `liste` de la fonction `uniformisation` désigne le même emplacement mémoire, leur _contenu_ (ce qui se trouve à cet emplacement mémoire) est donc le même; qu'on accède à cet emplacement par l'un ou l'autre est équivalent. En revanche on note que si on modifie la _valeur_ de `liste` (en écrivant `liste = ...`), cela n'a pas de conséquence sur `L`.

Notez bien la différence entre la fonction `uniformisation` ci-dessus et la fonction `uniformisation_ratee` ci-dessous :
```python
"""tentative d'uniformisation qui ne fonctionne pas"""
def uniformisation_ratee(liste):
    for e in liste :
        e = liste[0]
    
L = [1, 2, 3, 4, 5]
uniformisation_ratee(L)
```
sur [python tutor](https://pythontutor.com/visualize.html#code=%22%22%22tentative%20d'uniformisation%20qui%20ne%20fonctionne%20pas%22%22%22%0Adef%20uniformisation_ratee%28liste%29%3A%0A%20%20%20%20for%20e%20in%20liste%20%3A%0A%20%20%20%20%20%20%20%20e%20%3D%20liste%5B0%5D%0A%20%20%20%20%0AL%20%3D%20%5B1,%202,%203,%204,%205%5D%0Auniformisation_ratee%28L%29&cumulative=false&curInstr=0&heapPrimitives=nevernest&mode=display&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false). Ici la variable `e` désigne alternativement la valeur d'une case de la liste, puis la valeur de la première case de la liste, sans aucune conséquence sur le contenu de la liste elle-même.

### Exo 1 : Translation
Écrire et documenter une fonction `translate` qui prend en argument un point et un vecteur (dans cet ordre), chacun sous la forme d'une liste d'entiers à deux cases (sans avoir à vérifier les types) et transforme le point en lui faisant subir une translation du vecteur donné (cette fonction ne doit pas fabriquer une nouvelle liste pour un nouveau point, mais modifier le contenu de celle qui lui est transmise en argument).

<!-- grade_id: cell-a19867e9793ad84f, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-da90a68ff83a607f, points: 1, grade: true, solution: false -->
```python
P = [0, 0]
translate(P, [1, -2])
assert(P == [1, -2])

translate(P, [-5, 6])
assert(P == [-4, 4])
```

### Exo 2 : Annulation des impairs
Écrire et documenter une fonction `annulation_impairs` qui prend en argument une liste supposés composés d'entiers (sans avoir à le vérifier) et met à 0 chaque case qui contient un nombre impair (cette fonction ne doit pas fabriquer une nouvelle liste, mais modifier le contenu de celle qui lui est transmise en argument).

<!-- grade_id: cell-a1d851d03345f16e, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-d5006c90811e2e8e, points: 1, grade: true, solution: false -->
```python
L = []; annulation_impairs(L); assert(L == [])
L = [1]; annulation_impairs(L); assert(L == [0])
L = [2]; annulation_impairs(L); assert(L == [2])
L = [2, -1, 5, 6, 8, 10, 13]; annulation_impairs(L); assert(L == [2, 0, 0, 6, 8, 10, 0])
```

### Exo 3 : Permutation circulaire
Écrire et documenter une fonction `permutation_circulaire` qui prend en argument une liste et lui applique _en place_ (c'est-à-dire sans créer de nouvelle liste) une permutation circulaire gauche, c'est-à-dire que chaque élément est déplacé dans la case immédiatement précédente, sauf le premier qui est déplacé dans la dernière case.

Note pour désigner la dernière case d'une liste `L`, on a déjà vu qu'on peut écrire `L[len(L)-1]`, mais on peut aussi écrire `L[-1]`.

<!-- grade_id: cell-f23d09d5a0ab4d56, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-5052b0a04007b989, points: 1, grade: true, solution: false -->
```python
L = []
permutation_circulaire(L)
assert(L == [])

L = [1]
permutation_circulaire(L)
assert(L == [1])

L = [1, 2]
permutation_circulaire(L)
assert(L == [2, 1])

L = [5, 6, 7, 8]
permutation_circulaire(L)
assert(L == [6, 7, 8, 5])
```

### Exo 4 : Tri bulle
La problématique du tri est une problématique récurrente en informatique car bien des problèmes sont simplifiés si les données sont triées (pensez à ce que serait la recherche dans un dictionnaire si les mots apparaissaient dans un ordre quelconque). Il est donc souvent préférable de commencer par trier les données, et ensuite travailler sur ce nouvel objet que de travailler sur un ensemble non structuré de données.

Il existe beaucoup d'algorithmes classiques de tri. Nous allons étudier le _tri bulle_ qui est un algorithme facile à mettre en œuvre, mais pas très efficace (il est lent), sauf si les données de départ sont peu désordonnées.

Le principe de base du tri bulle est de parcourir les données de gauche à droite en comparant chaque donnée à sa voisine : si elle ne sont pas dans le bon ordre, on les permute et on continue, sinon on continue. À la fin de ce parcours (qu'on appelle un _passage_), les données ne sont a priori toujours pas triée, mais elles sont moins désordonnées, et en particulier le plus grand élément se trouve à la fin.

Un exemple de **premier passage** :

$\begin{array}{|c|c|c|c|c|}
\hline
\stackrel{\bullet}{4} & \stackrel{\bullet}{2} & 5 & 1 & 3\\ \hline
\end{array}$ 4 et 2 ne sont pas dans le bon ordre $\rightarrow$ on les permute et on continue

$\begin{array}{|c|c|c|c|c|}
\hline
2 & \stackrel{\bullet}{4} & \stackrel{\bullet}{5} & 1 & 3\\ \hline
\end{array}$ 4 et 5 sont dans le bon ordre $\rightarrow$ on continue

$\begin{array}{|c|c|c|c|c|}
\hline
2 & 4 & \stackrel{\bullet}{5} & \stackrel{\bullet}{1} & 3\\ \hline
\end{array}$ 5 et 1 ne sont pas dans le bon ordre $\rightarrow$ on les permute et on continue

$\begin{array}{|c|c|c|c|c|}
\hline
2 & 4 & 1 & \stackrel{\bullet}{5} & \stackrel{\bullet}{3}\\ \hline
\end{array}$ 5 et 3 ne sont pas dans le bon ordre $\rightarrow$ on les permute et on a fini le passage

$\begin{array}{|c|c|c|c|c|}
\hline
2 & 4 & 1 & 3 & 5\\ \hline
\end{array}$

On recommence ensuite ces opérations, mais sans traiter la dernière case, puisqu'elle contient déjà le plus grand élément :

**Deuxième passage** :

$\begin{array}{|c|c|c|c|c|}
\hline
\stackrel{\bullet}{2} & \stackrel{\bullet}{4} & 1 & 3 & {\color{gray} 5}\\ \hline
\end{array}$ 2 et 4 sont dans le bon ordre $\rightarrow$ on continue

$\begin{array}{|c|c|c|c|c|}
\hline
2 & \stackrel{\bullet}{4} & \stackrel{\bullet}{1} & 3 & {\color{gray} 5}\\ \hline
\end{array}$ 4 et 1 ne sont pas dans le bon ordre $\rightarrow$ on les permute et on continue

$\begin{array}{|c|c|c|c|c|}
\hline
2 & 1 & \stackrel{\bullet}{4} & \stackrel{\bullet}{3} & {\color{gray} 5}\\ \hline
\end{array}$ 4 et 3 ne sont pas dans le bon ordre $\rightarrow$ on les permute et on a fini le passage

$\begin{array}{|c|c|c|c|c|}
\hline
2 & 1 & 3 & 4 & {\color{gray} 5}\\ \hline
\end{array}$

**Troisième passage** :

$\begin{array}{|c|c|c|c|c|}
\hline
\stackrel{\bullet}{2} & \stackrel{\bullet}{1} & 3 & {\color{gray} 4} & {\color{gray} 5}\\ \hline
\end{array}$ 2 et 1 ne sont pas dans le bon ordre $\rightarrow$ on les permute et on continue

$\begin{array}{|c|c|c|c|c|}
\hline
1 & \stackrel{\bullet}{2} & \stackrel{\bullet}{3} & {\color{gray} 4} & {\color{gray} 5}\\ \hline
\end{array}$ 2 et 3 sont dans le bon ordre $\rightarrow$ on a fini le passage

$\begin{array}{|c|c|c|c|c|}
\hline
1 & 2 & 3 & {\color{gray} 4} & {\color{gray} 5}\\ \hline
\end{array}$

**Quatrième passage** :

$\begin{array}{|c|c|c|c|c|}
\hline
\stackrel{\bullet}{1} & \stackrel{\bullet}{2} & {\color{gray} 3} & {\color{gray} 4} & {\color{gray} 5}\\ \hline
\end{array}$ 1 et 2 sont dans le bon ordre $\rightarrow$ on a fini le passage

$\begin{array}{|c|c|c|c|c|}
\hline
1 & 2 & {\color{gray} 3} & {\color{gray} 4} & {\color{gray} 5}\\ \hline
\end{array}$

Comme le dernier passage ne concernait que deux éléments, c'est fini, les données sont triées. À noter que si on avait fait un passage sans aucune permutation, on aurait pu s'arrêter avant (car alors tous les éléments auraient été dans le bon ordre).

Écrire et commenter une fonction `tri_bulle` qui prend en argument une liste de nombres (sans avoir à le vérifier) et la trie en place en suivant l'algorithme de tri bulle.

Notez que pour inverser les valeurs de deux variables `a` et `b` en python, il suffit d'écrire
```python
a, b = b, a
```
(Nous verrons plus tard une explication de cette notation.)

<!-- grade_id: cell-3878300614380675, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-892d1486978beb14, points: 1, grade: true, solution: false -->
```python
L = [4, 2, 5, 1, 3]; tri_bulle(L); assert(L == [1, 2, 3, 4, 5])
L = []; tri_bulle(L); assert(L == [])
L = [0]; tri_bulle(L); assert(L == [0])
L = [-1, 0]; tri_bulle(L); assert(L == [-1, 0])
L = [1, 0]; tri_bulle(L); assert(L == [0, 1])
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* Les passages des paramètres dans les fonctions se font par valeur : la valeur d'un nombre est le nombre lui-même, mais la valeur d'une liste est son adresse, on peut donc modifier le contenu d'une liste dans une fonction.
* Retenir le principe de l'algorithme du tri bulle.
* On peut inverser les valeurs de `a` et `b` avec `a, b = b, a`.
* `for i in range(a, b) :` où `a` et `b` sont deux entiers permet de parcourir dans l'ordre croissant les entiers de `a` à `b`-1 (ensemble qui peut être vide suivant les valeurs de `a` et `b`).
* Le dernier élément d'une liste `L` peut être désigné par `L[-1]`.
</div>