# ITC - MPSI
---

# TP5 : Retour sur les listes
Ce TP revient sur les notions vues depuis le début de l'année et les complète. C'est pour vous l'occasion de vous rendre compte si vous les avez bien assimilées. Sinon n'hésitez pas à reprendre les anciens TP.

### Exo 1 : Calculatrice
Dans cet exercice, il faut implémenter une fonction `calculatrice` qui prend en argument une liste décrivant un calcul à faire entre entiers et qui renvoie le résultat (en supposant, sans avoir à le vérifier, que la liste est admissible).

Les listes admissibles et leurs évaluations sont les suivantes :
* `[n]`, où `n` est un entier, est une liste admissible; son évaluation est la valeur de `n`<br />
exemple : `calculatrice([5])` renvoie 5
* `[L1, '+', L2]`, où `L1` et `L2` sont des listes admissibles, est une liste admissible; son évaluation est `e1 + e2`, où `e1` est l'évaluation de `L1` et `e2` celle de `L2`<br />
exemples :
  * `calculatrice([[5], '+', [-1]])` renvoie 4
  * `calculatrice([[3], '+', [[5], '+', [-1]] ])` renvoie 7
  
* de même pour `[L1, '-', L2]` et `[L1, '*', L2]`
* `[L1, '/', L2]`, où `L1` et `L2` sont des listes admissibles, est une liste admissible; son évaluation est `e1 // e2`, où `e1` est l'évaluation de `L1` et `e2` celle de `L2`<br />
    exemple : `calculatrice([[3], '/', [[5], '+', [-1]] ])` renvoie 0

1. Écrire une fonction `eval` qui prend trois arguments : un nombre, un caractère dont on suppose (sans avoir à le vérifier) que c'est `+`, `-`, `*` ou `/`, et un autre nombre, et renvoie le résultat de l'opération correspondante entre entiers.

<!-- grade_id: cell-b569a17d50ad8b35, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-a09344605aa754e2, points: 1, grade: true, solution: false -->
```python
assert(eval(3, '*', 15) == 45)
assert(eval(3, '+', 15) == 18)
assert(eval(3, '-', 15) == -12)
assert(eval(6, '/', 4) == 1)
```

2. Écrire et documenter la fonction `calculatrice`. Remarquez qu'une liste admissible est soit de longueur 1 et dans ce cas le contenu de cette case fournit immédiatement son évaluation, soit de longueur 3, et dans ce cas il faut évaluer les cases 0 et 2, puis faire la bonne opération entre les deux.

<!-- grade_id: cell-a7e5ba15c8b50cde, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-ddffc9dac4e5de8d, points: 1, grade: true, solution: false -->
```python
assert(calculatrice([5]) == 5)
assert(calculatrice([[5], '+', [-1]]) == 4)
assert(calculatrice([[3], '+', [[5], '+', [-1]] ]) == 7)
assert(calculatrice([[3], '/', [[5], '+', [-1]] ]) == 0)
```

## Ajout et suppression d'éléments dans une liste
Nous avons vu qu'une liste en `python` se note entre crochets : `[1, 2, 3, -4]`. Jusqu'à présent, une fois une liste créée, les seules modifications que nous avons faites sont la modification de la valeur d'une case de cette liste :

```python
L = ['a', 'b', 'c', 'd']
print(str(L) + '\t# valeur initiale de L')

L[0] = 'z'
print(str(L) + '\t# L après la modification de la case 0')

```

En fait, on peut ajouter un élément à la fin d'une liste en utilisant une _méthode_ (le fait de comprendre ce qu'est une méthode et comment en écrire une est hors programme, mais il faut savoir utiliser une méthode qui vous est fournie). Une méthode se comporte comme une fonction, mais la syntaxe pour l'invoquer est différente. Pour ajouter un élément à la fin d'une liste, on utilise la méthode `append` :

```python
L = ['a', 'b', 'c', 'd']

print(str(L) + '\t\t# valeur initiale de L')

L.append(2)  # observezez ici : nom de la liste, caractère point, nom de la méthode, 
             #                  arguments pour la méthode entre parenthèses
    
print(str(L) + "\t\t# L apres l'ajout de 2")

L.append('toto')

print(str(L) + "\t# L apres l'ajout de 'toto'")
```

On peut aussi suprrimer le dernier élément d'une liste, tout en récupérant sa valeur grâce à la méthode `pop` :

```python
L = ['a', 'b', 'c', 'd']

print(str(L) + '\t# valeur initiale de L')

n = L.pop()  # observez ici : il faut des parenthèses même en l'absence d'argument

print(n + '\t\t\t# valeur renvoyée par pop')
print(str(L) + "\t\t# L après l'appel à pop")
```

Attention à ce qui se passe si on essaie de faire un `pop` sur une liste vide :

```python
[].pop()
```

Il faut donc toujours tester si une liste est non vide (par exemple en utilisant la fonction `len` qui donne sa longueur) avant d'essayer de récupérer un élément de la liste.

(À noter que les programmeurs `python` ne le font pas comme ça, mais utilise `try... catch` qui est hors programme. Cette remarque est donc là uniquement pour votre culture.)

Le comportement d'une liste avec les méthodes `append` et `pop` est le comportement d'une _pile_ : on ajoute des éléments et on récupère systématiquement l'élément le plus récent qui appartient à la pile. Cette structure de données est une structure de données classique en informatique et elle apparaît dans de nombreux algorithmes (un exemple d'utilisation à l'exercice 4 ci-dessous).

Le vocabulaire associé à ces opérations est _empiler_ quand on ajoute un nouvel élément et _dépiler_ quand on retire l'élément le plus récent de la pile.

## Tester l'appartenance à une liste
Vous avez écrit dans un TP une fonction qui teste si un élément appartient à une liste. Cette fonctionnalité est en fait présente nativement en `python`, grâce à l'opérateur `in` : `x in L`, où `x` est une expression (par exemple une valeur ou une variable) et `L` une liste, vaut `True` si et seulement si l'évaluation de `x` est un élément de `L`.

```python
'a' in L
```

```python
L = [2, 'a']
1+1 in L
```

```python
'b' in L
```

### Exo 2 : Concaténation de listes
Écrire et documenter une fonction `concat` qui prend en argument deux listes et modifie la première en lui ajoutant le contenu de la deuxième à la fin (dans l'ordre), en vous servant de la méthode `append`.

<!-- grade_id: cell-8009abb96b81614b, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-dca253d7919dfc43, points: 1, grade: true, solution: false -->
```python
L = []
concat(L, [1, 4, 7, -2])
assert(L == [1, 4, 7, -2])

concat(L, [])
assert(L == [1, 4, 7, -2])

L2 = ['b', 'n']
concat(L, L2)
assert(L == [1, 4, 7, -2, 'b', 'n'])
assert(L2 == ['b', 'n'])
```

### Exo 3 : Extraction des consonnes
Écrire et documenter une fonction `consonantique` qui prend en argument une liste de chaînes de caractères de longueur 1 qui sont toutes des lettres minuscules sans accents ou des espaces (sans avoir à le vérifier) et renvoie une nouvelle liste de chaînes de caractères de longueur 1 contenant uniquement les consonnes de la liste d'origine (dans le même ordre).

<!-- grade_id: cell-cd5ad0a4a75f3a86, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-50bb483547b3e1ac, points: 1, grade: true, solution: false -->
```python
assert(consonantique([]) == [])

assert(consonantique(['a']) == [])

assert(consonantique(['b']) == ['b'])

phrase = list("ceci est une phrase pour faire un test")
assert(consonantique(phrase) == ['c', 'c', ' ', 's', 't', ' ', 'n', ' ', 'p', 'h', 'r', 's', ' ', 
                    'p', 'r', ' ', 'f', 'r', ' ', 'n', ' ', 't', 's', 't'])
assert(phrase == list("ceci est une phrase pour faire un test"))
```

### Exo 4 : Une calculatrice plus sophistiquée
On désire maintenant écrire une fonction `calculatrice_scientifique` qui prend en argument une liste représentant une expression supposée bien formée (sans avoir à le vérifier) et renvoie le résultat du calcul entre entiers.

Une expression bien formée est composée de la manière suivante :
* $n$, où $n$ est un entier, est une expression bien formée, son évaluation est la valeur de $n$<br />
exemple : 5 est une expression bien formée, elle s'évalue en $5$
* $(e_1 + e_2)$, où $e_1$ et $e_2$ sont des expressions bien formées, est une expression bien formée, elle s'évalue en la somme des évaluations de $e_1$ et $e_2$<br />
exemple :
  * $(5 + -1)$ est une expression bien formée, elle s'évalue en $4$
  * $(3 + (5 + -1))$ est une expression bien formée, elle s'évalue en $7$
* de même pour $(e_1 - e_2)$ et $(e_1 * e_2)$
* $(e_1 / e_2)$, où $e_1$ et $e_2$ sont des expressions bien formées, est une expression bien formée, elle s'évalue en le quotient entier des évaluations de $e_1$ et $e_2$<br />
exemples :
  * $(5 / -1)$ est une expression bien formée, elle s'évalue en $-5$
  * $(3 / (5 + -1))$ est une expression bien formée, elle s'évalue en $0$

À partir d'une expression bien formée, on obtient une liste contenant des entiers et les caractères '(', ')', '+', '-', '*' et '/' assez naturellement.

exemples:
* `[ -5, '*', '(', 13, '+', '(', 6, '/', '(', 1, '+', 2, ')', ')', ')' ]` est une liste admissible pour la fonction `calculatrice_scientifique`, elle s'évalue en -75 car elle correspond au calcul `-5*(13+(6//(1+2)))` en `python`
* `['-', 5]` n'est pas une liste admissible (car on n'autorise pas le `-` unaire`)
* `['(', 3, '+', 2 ]` n'est pas une liste admissible (car elle est mal parenthésée)

L'algorithme à suivre pour implémenter une telle calculatrice est le suivant :
* si la liste fournie est de longueur 1, on renvoie tout de suite le résultat
* on ajoute une parenthèse droite à la liste fournie
* on crée une pile vide `operateurs`
* on crée une pile vide `nombres`
* on parcourt la liste donnée en argument de gauche à droite :
  * si le caractère courant est une parenthèse gauche, on l'ignore,
  * sinon :
    * si c'est un opérateur, on l'empile dans `operateurs`
    * si c'est un nombre, on l'empile dans `nombres`
    * si c'est une parenthèse droite, on dépile les deux derniers nombres et le dernier opérateur, et on empile le résultat de l'opération correspondante dans `nombres`
* on supprime la parenthèse qu'on avait ajoutée à la liste transmise en argument
* le résultat attendu est au sommet de la pile `nombres`


Écrire la fonction `calculatrice_scientifique` (vous pouvez et même devez s'il y a besoin réutiliser la fonction `eval` de l'exercice 1).

N'hésitez pas à faire du pas à pas dans [python tutor](https://pythontutor.com/) pour vous aider.

<!-- grade_id: cell-df37fc4c5e693bf1, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-d72806137cbd9e50, points: 2, grade: true, solution: false -->
```python
assert(calculatrice_scientifique([5]) == 5)

L = [-5, '*', '(', 13, '+', '(', 6, '/', '(', 1, '+', 2, ')', ')', ')']
L2 = L[:] # pour faire une copie de la liste
assert(calculatrice_scientifique(L) == -75)
assert(L == L2)
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* syntaxe pour utiliser une méthode
* méthode `append` pour ajouter un élément à la fin d'une liste
* méthode `pop` pour récupérer et supprimer le dernier élément d'une liste non vide
* erreur de `pop` si la liste est vide
* opérateur `in` pour tester l'appartenance d'un élément à une liste
* structure de pile implémentée par une liste / empiler / dépiler
</div>