# ITC - MPSI
---

# TP6 : Récursivité et plongée dans la nature des listes
On a vu dans le TP5 un exemple de fonction qui s'appelle elle-même. C'est ce qu'on appelle une fonction _récursive_.

La démarche quand on implémente une fonction récursive est très proche de la démarche lors d'une démonstration par récurrence:
* on sait résoudre le problème pour le cas de base ou pour des cas de base,
* on suppose qu'on sait le faire à un certain rang ou jusqu'à un certain rang, et alors on le résoud pour le rang suivant en se ramenant à des instances plus petites du même problème.

#### Exemple : de combien de façon peut-on monter un escalier à $n$ marches ?
Supposons qu'on puisse monter les marches soit une à la fois, soit deux à la fois et qu'on se trouve devant un escalier comportant $n$ marches, comment calculer le nombre de façons de monter cet escalier ? Notons $F_n$ ce nombre.

**case de base** : 
* si $n=0$, il y a une seule façon de monter l'escalier : on reste là où on est
* si $n=1$, il y a une seule façon de monter l'escalier : on franchit l'unique marche

**héréditée** : si on suppose qu'on sait compter le nombre de façons de monter un escalier ayant au plus $n$ marches pour un certain $n\geq 1$, alors pour un escalier ayant $n+1$ marche :
* soit on commence par franchir un seule marche et on se retrouve devant un escalier de $n$, on a donc $F_n$ façons de franchir les marches restantes,
* soit on commence par enjamber deux marches d'un coup et on se retrouve devant un escalier de $n-1$ marches, on a donc $F_{n+1}$ façons de franchir les marches restantes.

La fonction associée est donc :

```python
"""n : un entier positif ou nul qui représente le nombre de marches d'un escalier
sortie : le nombre de façons de monter l'escalier avec des pas d'une ou deux marches"""
def montee_escaliers(n):
    if n == 0 or n == 1:
        return 1
    return montee_escaliersiers(n-1) + montee_escaliers(n-2)
```

Si on n'a pas de cas de base on risque de faire des appels récursifs à l'infini, et comme la mémoire du processus n'est pas illimitée, on risque un débordement de pile.

Vous pouvez regarder un exemple de débordement de pile sur python tutor (ici embarqué dans le notebook) :

```python
from IPython.display import IFrame
IFrame('https://pythontutor.com/iframe-embed.html#code=def%20f%28%29%20%3A%0A%20%20%20%20f%28%29%0A%20%20%20%20%0Af%28%29&codeDivHeight=400&codeDivWidth=350&cumulative=false&curInstr=0&heapPrimitives=nevernest&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false', width=980, height=350)
```

### Exo 1 : Calcul récursif de puissance
En remarquant que $x^{2k}=(x^k)^2$ et $x^{2k+1}=x(x^k)^2$, écrire et documenter une fonction **récursive** `puissance` qui permet de calculer une puissance entière positive d'un entier (sans utiliser l'opérateur python `**`).

<!-- grade_id: cell-757a46694f22026b, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-af35c7281e8ce885, points: 2, grade: true, solution: false -->
```python
assert(puissance(5, 0) == 1)
assert(puissance(2, 12) == 2**12)
```

## C'est quoi la valeur d'une liste?
On a vu qu'une variable sert à désigner une valeur. Ainsi, après ces instructions :
```python
p = 2.5
q = p
```
les variables `p` et `q` désignent la même valeur, comme on peut le voir ici :

```python
from IPython.display import IFrame
IFrame('https://pythontutor.com/visualize.html#code=p%20%3D%202.5%0Aq%20%3D%20p&cumulative=false&curInstr=0&heapPrimitives=true&mode=display&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false', width=980, height=350)
```

Si on change par la suite la valeur désignée par `p`, comme ici :
```python
p = 2.5
q = p
p = 1.
```
on voit que `q` continue à désigner la même valeur, et c'est `p` qui change de cible.

```python
from IPython.display import IFrame
IFrame('https://pythontutor.com/visualize.html#code=p%20%3D%202.5%0Aq%20%3D%20p%0Ap%20%3D%201.&cumulative=false&curInstr=0&heapPrimitives=true&mode=display&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false', width=980, height=350)
```

Que se passe-t-il quand une variable désigne une liste ?
```python
L = [1, 2, 3]
M = L
N = L
N = [1, 2]
M[1] = 'a'
```
Regardons graphiquement :

```python
from IPython.display import IFrame
IFrame("https://pythontutor.com/visualize.html#code=L%20%3D%20%5B1,%202,%203%5D%0AM%20%3D%20L%0AN%20%3D%20L%0AN%20%3D%20%5B1,%202%5D%0AM%5B1%5D%20%3D%20'a'&cumulative=false&curInstr=0&heapPrimitives=nevernest&mode=display&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false", width=980, height=350)
```

Si on modifie la liste désignée par `N`, cela n'a pas de conséquence sur la liste désignée par `L`, alors que si on modifie une case **dans la liste** désignée par `M`, la liste désignée par `L` étant la même, ce changement aura des conséquences.

Ainsi si dans une fonction récursive on transmet la même liste d'appel en appel, les modifications apportées à cette liste à chaque appel ont des répercussions pour tous les appels en cours.

Regardons par exemple le code suivant:

```python
from IPython.display import IFrame
IFrame("https://pythontutor.com/visualize.html#code=%22%22%22L%20%3A%20liste%20d'entiers%0An%20%3A%20entier%20inf%C3%A9rieur%20%C3%A0%20len%28L%29%0Asortie%20%3A%20tous%20les%20%C3%A9l%C3%A9ments%20impairs%20de%20L%20d'indice%20au%20moins%20n%20on%20%C3%A9t%C3%A9%20mis%20%C3%A0%20z%C3%A9ro%22%22%22%0Adef%20annuler_les_impairs%28L,%20n%29%3A%0A%20%20%20%20if%20L%5Bn%5D%252%20%3D%3D%201%3A%0A%20%20%20%20%20%20%20%20L%5Bn%5D%20%3D%200%0A%20%20%20%20if%20n%2B1%20%3C%20len%28L%29%3A%0A%20%20%20%20%20%20%20%20annuler_les_impairs%28L,%20n%2B1%29%0A%0AL%20%3D%20%5B1,%202,%203,%204%20,%205,%206,%207%5D%0Aannuler_les_impairs%28L,%200%29&cumulative=false&curInstr=0&heapPrimitives=nevernest&mode=display&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false", width=980, height=550)
```

### Exo 2 : Croissance
Écrire et documenter une fonction **récursive** `fin_croissante` qui prend en argument une liste de nombres et un indice dans cette liste (sans avoir à le vérifier) et teste si les nombres de la liste sont dans l'ordre croissant (au sens large) à partir de l'indice.

<!-- grade_id: cell-db7072e3ef90dcda, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-b728248fe704c93f, points: 2, grade: true, solution: false -->
```python
assert(fin_croissante([3, 3, 3, 2, 3, 4, 4], 3))
assert(not fin_croissante([3, 3, 3, 2, 3, 4, 4], 2))
assert(fin_croissante([1], 0))
```

### Exo 3 : Remplissage par diffusion 1D
Le _remplissage par diffusion_ est un algorithme d'infographie qui permet de colorer tous les pixels d'une forme fermée d'une même couleur (sous `gimp`, pour ceux qui connaissent, cela correspond à l'icône du petit seau qui se renverse).

Nous allons appliquer ce principe en dimension 1. On dispose d'un tableau d'entier, et d'une position dans ce tableau, et il s'agit de mettre à 0 toutes les cases ayant la même valeur que la case à la position donnée et atteignables en ne passant que par des cases qui ont cette même valeur.

exemple :

tableau de départ : $\begin{array}{|c|c|c|c|c|c|c|c|} \hline 2 & 3 & 1 & 1 & 1 & 2 & 2 & 1\\ \hline\end{array}$
, position : 2

tableau d'arrivée : $\begin{array}{|c|c|c|c|c|c|c|c|} \hline 2 & 3 & 0 & 0 & 0 & 2 & 2 & 1\\ \hline\end{array}$

Écrire et documenter une fonction **récursive** `remplissage1D` qui réalise cette opération.

<!-- grade_id: cell-02d3806600ef2806, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-c348d1fac3d6b1a7, points: 2, grade: true, solution: false -->
```python
tab = [2, 3, 1, 1, 1, 2, 2, 1]
remplissage1D(tab, 2)
assert(tab == [2, 3, 0, 0, 0, 2, 2, 1])

remplissage1D(tab, 0)
assert(tab == [0, 3, 0, 0, 0, 2, 2, 1])

tab = [2, 3, 1, 1, 1, 2, 2, 1, 1, 1, 1]
remplissage1D(tab, 8)
assert(tab == [2, 3, 1, 1, 1, 2, 2, 0, 0, 0, 0])
```

<!-- grade_id: cell-c348d1fac3d6b1a8, points: 1, grade: true, solution: false -->
```python
tab = [2, 3, 1, 1, 1, 0, 0, 1, 1, 1, 1]
remplissage1D(tab, 5)
assert(tab == [2, 3, 1, 1, 1, 0, 0, 1, 1, 1, 1])
```

## Listes multi-dimensionnelles
Jusqu'à présent, on a surtout utilisé des types numériques dans les cases des listes, mais on peut également utiliser une liste comme élément d'une liste.

Par exemple :
```python
L = [ [1, 2, 3], [4, 5, 6] ]
```
`L[0]` désigne alors la liste de la case 0, soit la liste `[1, 2, 3]` et `L[1]` la liste de la case 1, soit `[4, 5, 6]`.

### Exo 4 : Rectangle
Écrire et documenter une fonction `rectangle` qui prend en argument une liste de listes et vérifient que les listes de toutes les cases ont la même longueur.

<!-- grade_id: cell-891960d07a1e1dc2, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-eab87d2449f077f4, points: 2, grade: true, solution: false -->
```python
assert(rectangle([]))
assert(rectangle([[], [], [], []]))
assert(not rectangle([[], [], [], [1]]))
assert(not rectangle([[1], [1], [1, 2], [5]]))
assert(rectangle([[1], [2], [3], [4]]))
assert(rectangle([[4, 5, -3, 6]]*10))
```

### Exo 5 : Remplissage par diffusion 2D
Le principe est le même en 2D qu'en 1D, mais chaque point a maintenant 4 voisins.

Écrire et documenter une fonction **récursive** `remplissage2D` qui réalise cette opération : on lui donne une liste de listes d'entiers supposées toutes de la même longueur (sans avoir à le vérifier) et deux entiers représentant des coordonnées dans le tableau, et la fonction met à 0 toutes les cases ayant la même valeur que la case à la position donnée et atteignables en ne passant que par des cases qui ont cette même valeur.

<!-- grade_id: cell-8bf52cadd9fe2603, points: 2, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

Pour tester votre code, exécuter les trois cellules suivantes :

```python
import random
from IPython.display import Image

"""nom : nom d'un fichier
image : liste de listes d'entiers positifs ou nuls
couleurs : tableau de couleurs
sortie : fichier image d'extension ppm créé et représentant l'image"""
def creer_image(nom, image, couleurs = None):
    nb_couleurs = max( [ max(ligne) for ligne in image ] )
    if not couleurs:
        couleurs = [ None ] * (nb_couleurs+1)
        couleurs[0] = (random.randint(0, 255), random.randint(0, 255), random.randint(0, 255))
    with open(nom+'.ppm', 'w') as fichier:
        fichier.write('P3\n')
        fichier.write(str(len(image[0])) + ' ' + str(len(image)) + '\n')
        fichier.write('256\n')
        for ligne in image:
            for pixel in ligne:
                if not couleurs[pixel]:
                    couleurs[pixel] = (random.randint(0, 255), random.randint(0, 255), random.randint(0, 255))
                fichier.write(str(couleurs[pixel][0]) + ' '
                             + str(couleurs[pixel][1]) + ' '
                             + str(couleurs[pixel][2]) + ' ')
            fichier.write('\n')
    fichier.close()
    return couleurs

n = 15
k = 30
image = [ [1]*(k+2*n) for i in range(n) ]
image.append( [1]*n+[2]*k+[1]*n )
for i in range(2*n):
    image.append( [1]*n + [2] + [1]*(k-2) + [2] + [1]*n )
image.append( [1]*n+[2]*k+[1]*n )
image.extend([ [1]*(k+2*n) for i in range(n) ])

couleurs = creer_image('toto', image)

!convert toto.ppm toto.png  # pas du python; crée un fichier d'extension png
!rm toto.ppm                # pas du python; supprime le fichier d'extension ppm

print("image d'origine :")
Image('toto.png')
```

```python
remplissage2D(image, len(image[0])//2, len(image)//2)
creer_image('toto2', image, couleurs)

!convert toto2.ppm toto2.png  # pas du python; crée un fichier d'extension png
!rm toto2.ppm                 # pas du python; supprime le fichier d'extension ppm

print("image remplie à partir du centre :")
Image('toto2.png')
```

```python
n = 15
k = 30
image = [ [1]*(k+2*n) for i in range(n) ]
image.append( [1]*n+[2]*k+[1]*n )
for i in range(2*n):
    image.append( [1]*n + [2] + [1]*(k-2) + [2] + [1]*n )
image.append( [1]*n+[2]*k+[1]*n )
image.extend([ [1]*(k+2*n) for i in range(n) ])

remplissage2D(image, 0, 0)
creer_image('toto3', image, couleurs)

!convert toto3.ppm toto3.png  # pas du python; crée un fichier d'extension png
!rm toto3.ppm                 # pas du python; supprime le fichier d'extension ppm

print("image remplie à partir d'un coin :")
Image('toto3.png')
```