# ITC - MPSI
---

# TP7 : Retour sur la récursion

Faisons un parallèle (encore) entre fonctions récursives et démonstration par récurrence :

|| démonstration par récurrence|fonctions récursives|
|--|--|--|
|but| montrer que $P(n)$ est vrai pour tout $n\geq 0$ | calculer $f(n)$ pour tout $n\geq 0$|
|initialisation | montrer que $P(0)$ est vrai | savoir calculer $f(0)$|
|hypothèse| $P(k)$ vrai pour (ou jusqu'à) un certain rang $n-1$ | on sait calculer $f(k)$ jusqu'à un certain rang $n-1$|
|soit hérédité 1| montrer que $P(n-1)\Rightarrow P(n)$ | exprimer $f(n)$ en fonction de $f(n-1)$|
|soit hérédité 2| montrer que $\left(\forall k<n, P(n-1)\right)\Rightarrow P(n)$ | exprimer $f(n)$ en fonction de certains $f(k)$ pour $k<n$|
|conclusion | $P(n)$ vrai pour tout $n$ | on a une méthode effective pour calculer $f(n)$ tout $n$|

Une fonction récursive suit donc le schéma suivant :
* si on est dans le cas de base (=l'initialisation), on renvoie le résultat
* sinon on calcule récursivement (c'est-à-dire en faisant appel à la fonction qu'on est en train d'écrire) les valeurs pour les paramètres plus petits qui sont nécessaires pour trouver le résultat, on les combine et on renvoie le résultat.

#### Mais comment ça marche ?
Ça marche comme pour la récurrence. Pourquoi la conclusion de la récurrence est-elle vraie ? Supposons $N\in{\mathbb N}$. $P(0)$ est vraie et pour tout $n>0$, $P(n-1)\Rightarrow P(n)$, donc $P(1)$ est vraie. De la même façon, $P(2)$ est vraie, etc. jusqu'à arriver à $P(N)$.

Dans le cas des fonctions récursives: on sait calculer $f(0)$ et on sait exprimer $f(n)$ en fonction de $f(n-1)$ pour tout $n>0$, donc on sait calculer $f(1)$. De la même façon on sait calculer $f(2)$, etc. jusqu'à savoir calculer $f(N)$.

#### Oui mais comment la machine sait ce qu'il faut faire ?
Si on appelle $f$ avec le paramètre $N$, le code de $f$ va contenir un appel à $f$ avec un paramètre strictement plus petit, $N-1$. Cet appel va être lancé et il appellera lui-même $f$ avec un paramètre encore plus petit, $N-2$, etc. jusqu'à un appel avec le paramètre 0. Cet appel-là renvoie une réponse sans faire de nouvel appel. Réponse qui est transmise à l'appel qui avait le paramètre 1 et qui sera donc capable de construire la réponse pour 1 et de la transmettre à l'appel qui cherche à calculer la réponse pour 2, etc. et on finit par revenir à l'appel qui cherche à construire la réponse pour $N$ après avoir récupéré la réponse pour $N-1$.

Regardez ce qui se passe avec la fonction `factorielle` :

```python
from IPython.display import IFrame
IFrame("https://pythontutor.com/visualize.html#code=def%20factorielle%28n%29%3A%0A%20%20%20%20%22%22%22renvoie%20la%20factorielle%20de%20n,%20entier%20positif%20ou%20nul%22%22%22%0A%20%20%20%20assert%28n%20%3E%3D%200%29%0A%20%20%20%20if%20n%20%3D%3D%200%3A%0A%20%20%20%20%20%20%20%20return%201%0A%20%20%20%20return%20n%20*%20factorielle%28n-1%29%0A%20%20%20%20%0Afactorielle%283%29&cumulative=false&curInstr=0&heapPrimitives=nevernest&mode=display&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false", width=980, height=500)
```

### Exo 1 : Horloge
La fonction suivante permet d'afficher l'heure si $n$ a une valeur entre 1 et 12. Modifiez-la en ajoutant un appel récursif pour qu'elle donne l'heure pour $n$ entre 1 et 24:
```python
def horloge(n):
    """affiche l'horloge correspondant à l'heure n, pour n entre 1 et 12"""
    assert(n>=1 and n<=12)
    print(chr(128335+n))
```

<!-- grade_id: cell-17c3e7bf70539aba, points: 1, grade: true, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

### Exo 2 : pgcd
On se propose de calculer le pgcd de deux entiers positifs. Pour cela, on part de la remarque suivante : soit deux entiers strictement positifs $m$ et $n$, avec $m>n$, un entier $d$ divise $m$ et $n$ si et seulement s'il divise $n$ et $m-n$. D'un autre côté, le pgcd d'un nombre $n>0$ et de 0 est $n$.

Écrire et documenter une fonction **récursive** `pgcd` qui prend en argument deux entiers positifs et renvoie leur pgcd.

Vous pouvez obtenir le maximum entre deux nombres avec la fonction `max` et le minimum avec la fonction `min` (qui sont toutes les deux disponibles de base en `python`).

<!-- grade_id: cell-3aac13437eb5e0f2, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-328a048065315abe, points: 2, grade: true, solution: false -->
```python
assert(pgcd(6, 4) == 2)
assert(pgcd(42, 42) == 42)
assert(pgcd(60, 12) == 12)
assert(pgcd(60, 70) == 10)
```

### Exo 3 : Triangle de Pascal
Le triangle de pascal est une technique qui permet d'obtenir une généralisation de l'identité remarquable $(a+b)^2 = a^2 + 2ab + b^2$, en donnant les coefficients de $(a+b)^n$. Pour rappel, ces coefficients sont les coefficients binomiaux :
$$(a+b)^n = \sum\limits_{k=0}^{n} \left(\begin{array}{c}n\\k\end{array}\right) a^kb^{n-k}\:.$$

Le principe est le suivant : on part de la liste $L_0=\begin{array}{|c|}\hline  1 \\\hline\end{array}$ et on crée une nouvelle liste $L_1$ contenant une case de plus que $L_0$ telle que la première et la dernière case de cette liste valent toutes les deux 1, et la $i^e$ case de $L_1$ , pour toute autre valeur de $i$, vaut $L_0[i-1]+L_0[i]$. Et ainsi de suite.

On obtient ainsi : $L_1 = \begin{array}{|c|c|}\hline  1 & 1\\\hline\end{array}$, $L_2= \begin{array}{|c|c|c|}\hline  1 & 2 & 1\\\hline\end{array}$ et $L_3 = \begin{array}{|c|c|c|c|}\hline  1 & 3 & 3 & 1\\\hline\end{array}$.

1. Écrire et documenter une fonction récursive `ligne_suivante` qui prend en argument une liste d'entiers supposée être une ligne dans le triangle de Pascal (sans avoir à le vérifier) et renvoie la ligne suivante sous la forme d'une nouvelle liste.

<!-- grade_id: cell-1e3b288494f16417, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-05e13316b36a4903, points: 2, grade: true, solution: false -->
```python
L0 = [1]
L1 = ligne_suivante(L0)
assert(L1 == [1, 1])

L2 = ligne_suivante(L1)
assert(L2 == [1, 2, 1])

L3 = ligne_suivante(L2)
assert(L3 == [1, 3, 3, 1])

L5 = ligne_suivante(ligne_suivante(L3))
assert(L5 == [1, 5, 10, 10, 5, 1])
```

2. Écrire et documenter une fonction récursive `triangle_de_Pascal` qui prend en argument un entier $n\geq 0$ et renvoie une liste qui contient, sous forme de $n+1$ sous-listes, les $n+1$ premiers niveaux du triangle de Pascal. Si $n$ vaut 4, on s'attend à
```
print(triangle_de_Pascal(5))
[[1], [1, 1], [1, 2, 1], [1, 3, 3, 1], [1, 4, 6, 4, 1], [1, 5, 10, 10, 5, 1]]
```

<!-- grade_id: cell-6c5a30f762a2c59d, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-3ebd8d4bc118dcac, points: 2, grade: true, solution: false -->
```python
assert(triangle_de_Pascal(5) == [[1], [1, 1], [1, 2, 1], [1, 3, 3, 1], [1, 4, 6, 4, 1], [1, 5, 10, 10, 5, 1]])
assert(triangle_de_Pascal(0) == [[1]])
```

### Exo 4 : Tours de Hanoï
![](https://upload.wikimedia.org/wikipedia/commons/6/60/Tower_of_Hanoi_4.gif)

Les _tours de Hanoï_ est un jeu mathématique classique. On dispose de trois emplacements et de disques de diamètres deux à deux distincts qui sont empilés du plus grand au plus petit sur un des emplacements. Il faut déplacer ces disques sur un autre des emplacements, en respectant les règles suivantes :
* on ne peut pas poser un disque hors d'un emplacement,
* on ne peut déplacer qu'un seul disque à la fois,
* on ne peut pas poser un disque sur un disque plus petit.

Ce jeu possède une solution récursive naturelle, car si on sait déplacer une tour de hauteur $n-1$ en suivant les règles, il est très simple de déplacer une tour de hauteur $n$. Qui plus est, il est immédiat de déplacer une tour de hauteur 1.

Je vous fournis une fonction qui permet de visualiser l'état des tours :

```python
def plateau(n, total, symbole='__'):
    return ' '*(total-n) + symbole*n + ' '*(total-n)

def affichage(hanoi):
    """hanoi : liste de trois sous-listes d'entiers;
    chaque sous-liste représente les disques empilés, du plus grand au plus petit
    sortie : affichage graphique des tours"""
    h = max([ len(tour) for tour in hanoi ])-1
    lg = max([ max(tour) for tour in hanoi if len(tour)>0 ])
    while h >= 0:
        for tour in hanoi:
            if len(tour) > h:
                print(plateau(tour[h], lg), end='\t')
            else:
                print(plateau(0, lg), end='\t')
        print()
        h = h-1
    for tour in hanoi:
        print(plateau(1, lg, '[]'), end='\t')
    print('\n\n')

print('1er exemple :')
affichage([[5,3,1], [2], [8,7,6,4]])

print('2eme exemple :')
affichage([[5,3,1], [], [8,7,6,4,2]])
```

Écrire et documenter une fonction récursive `hanoi` qui prend en paramètre :
* l'état des tours sous forme d'une liste de trois sous-listes; chaque sous-liste contenant les tailles des disques de celui situé le plus bas vers celui situé le plus haut,
* le nombre de disques à déplacer,
* le numéro de l'emplacement (0, 1 ou 2) où se trouvent ces disques,
* le numéro de l'emplacement (0, 1 ou 2) où ces disques doivent aller.

Cette fonction doit modifier le contenu de son paramètre (donc aucune création de liste) et afficher l'état des tours à chaque déplacement.

On ne demande pas que la fonction vérifie que les règles de déplacement sont respectées, mais on demande qu'en cas de configuration initiale avec tous les disques correctement placés sur une seule tour, les mouvements de votre fonction se fassent effectivement selon les règles de déplacement.

<!-- grade_id: cell-81b6bf6325ef7818, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-2404320fdef58b44, points: 2, grade: true, solution: false -->
```python
H = [[4,3,2,1],[],[]]
hanoi(H, 2, 0, 2)
assert(H == [[4, 3], [], [2, 1]])
```

<!-- grade_id: cell-3e65cc65b2b8c01a, points: 1, grade: true, solution: false -->
```python
H = [[], [4,3,2,1],[]]
hanoi(H, 4, 1, 0)
assert(H == [[4, 3, 2, 1], [], []])
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* structure d'une fonction récursive
* `python` fournit des fonctions `max`et `min`
</div>