# ITC - MPSI
---

# TP9 : Chaînes de caractères
Nous explorons ici plus en détails les chaînes de caractères en `python`.

Pour rappel, une _chaîne de caractères_ est une suite finie de caractères. En python, les chaînes de caractères s'écrivent entre guillemets simples ou doubles: `'exemple'` ou `"exemple"`, ou entre triples guillements simples ou doubles, ce qui permet d'inclure des retours à la ligne:

`"""ceci est un
exemple"""`

`'''ceci en
est un autre'''`

On peut accéder à chaque caractère d'une chaîne de caractères, grâce à la notation `[]` (comme pour les cases d'une liste). Le première caractère se trouve à l'indice 0:

```python
mot = 'abcedf'
mot[0]
```

On peut parcourir les caractères d'une chaîne de caractères avec une boucle `for`:

```python
mot = 'abcdef'
for lettre in mot :
    print(lettre)
```

On accède au nombre de caractères d'une chaîne par la fonction `len`:

```python
len('abcdef')
```

```python
len('')
```

**On ne peut pas modifier un caractère dans une chaîne :**

```python
mot = 'abcdef'
mot[1] = 'z'
```

Le fait que les valeurs d'un type ne soient pas modifiables s'appelle l'_immuabilité_. Le type chaîne de caractères est donc immuable en `python`.

### Exo 1 : Trouver une lettre dans une chaîne
Écrire et documenter une fonction `appartient` qui prend en argument une première chaîne de longueur 1 (on vérifiera cette longueur avec `assert`) et une deuxième chaîne de longueur quelconque et vérifie si la première chaîne apparaît dans la deuxième en faisant un parcours explicite. (Dans la vraie vie, on peut écrire directement `mot in lettre`, mais ce n'est pas ce qui est attendu ici.)

<!-- grade_id: cell-46db2defd4c50ce9, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-0ea737c9a6d1d494, points: 2, grade: true, solution: false -->
```python
assert(not appartient('a', ''))
assert(appartient('a', 'a'))
assert(appartient('a', 'abcdef'))
assert(appartient('a', 'fedcba'))
assert(appartient('a', 'cbabcd'))
assert(not appartient('a', 'b'))
assert(not appartient('a', 'bcdef'))
```

### Exo 2 : Voyelles vs. consonnes
Écrire et documenter une fonction `comparaison_vc` qui prend en argument une chaîne de caratères et renvoie `True` si et seulement si la chaîne contient strictement plus de voyelles que de consonnes, `False` sinon. On supposera, sans avoir à le vérifier que la chaîne est composée exclusivement de lettres non accentuées.

<!-- grade_id: cell-f48771c6cbe491c7, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-c3ff5cbe50c4b891, points: 2, grade: true, solution: false -->
```python
assert(comparaison_vc('aaaabbb'))
assert(not comparaison_vc('aaaabbbb'))
assert(not comparaison_vc(''))
assert(comparaison_vc('abababa'))
assert(not comparaison_vc('ababab'))
```

### Exo 3 : Trouver une chaîne dans une autre à un emplacement précis
Écrire et documenter une fonction `apparait` qui prend en argument deux chaînes de caractères et un entier et vérifie si la première chaîne apparaît dans la deuxième à l'indice donné.

À noter que le "ou" entre booléens en `python` s'écrie `or`.

<!-- grade_id: cell-3e5218106f7cc1b7, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-175ad3b0244afd4d, points: 2, grade: true, solution: false -->
```python
assert(apparait('', 'abcd', 2))
assert(not apparait('', 'abcd', 4))
assert(apparait('abc', 'abcdef', 0))
assert(not apparait('abc', 'abcdef', 1))
assert(apparait('def', 'abcdef', 3))
assert(not apparait('def', 'abcdef', 2))
assert(not apparait('def', 'abcdef', 5))
```

### Exo 4 : Trouver une chaîne dans une autre
Écrire et documenter une fonction `apparait` qui prend en argument deux chaînes de caractères et un entier et vérifie si la première chaîne apparaît dans la deuxième'.

Vous remarquez qu'en `python`, on peut avoir plusieurs fonctions du même nom, à condition qu'elles ne possèdent pas le même nombre de paramètres.

<!-- grade_id: cell-b77ea53324568a68, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-2776e4041ff97966, points: 2, grade: true, solution: false -->
```python
assert(apparait('', ''))
assert(apparait('', 'abcdef'))
assert(apparait('abc', 'abcdef'))
assert(apparait('abc', 'abc'))
assert(apparait('abc', 'defabc'))
assert(apparait('abc', 'defabcdef'))
assert(apparait('abc', 'abcabc'))
assert(not apparait('abc', 'defab'))
```

### Exo 5 : Miroir

On peut créer une nouvelle chaîne de caractères en concaténant des chaînes existantes grâce à l'opérateur `+`:

```python
'abc' + 'def'
```

Écrire et documenter une fonction `miroir` qui prend en argument une chaîne de caractères et renvoie une chaîne correspondant à la lecture de droite à gauche de la première.

<!-- grade_id: cell-f9e981eb6abf5fba, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-2bb337270033490a, points: 2, grade: true, solution: false -->
```python
assert(miroir('') == '')
assert(miroir('a') == 'a')
assert(miroir('aaa') == 'aaa')
assert(miroir('abcdef') == 'fedcba')
assert(not miroir('abcdef') == 'abcdef')
```

### Exo 6 : bégaiement simple
On peut créer une chaîne de caractères en répétant le motif d'une chaîne existante avec l'opérateur `*`:

```python
'abc' * 5
```

Écrire et documenter une fonction `begaiement` qui prend en argument une chaîne de caractères et un entier et renvoie une nouvelle chaîne dans laquelle chaque lettre est répétée, comme dans l'exemple suivant:

```python
begaiement('abc', 3)
'aaabbbccc'
```

<!-- grade_id: cell-4c48d94d4ba57919, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-5e4d8e400367e04c, points: 2, grade: true, solution: false -->
```python
assert(begaiement('', 3) == '')
assert(begaiement('abc', 3) == 'aaabbbccc')
assert(begaiement('abc', 0) == '')
assert(begaiement('aaa', 1) == 'aaa')
```

### Exo 7 : bégaiement
Écrire et documenter une fonction `begaiement` qui prend en argument une chaîne de caractères et un entier et renvoie une nouvelle chaîne dans laquelle chaque consonne est répétée, comme dans l'exemple suivant:

```python
begaiement('abc', 3)
'abbbccc'
```

<!-- grade_id: cell-f851de74fd11c575, points: 0, grade: false, solution: true -->
```python
# Écrire votre code ici
raise NotImplementedError # effacer cette ligne une fois le code écrit
```

<!-- grade_id: cell-4d8c26fe25639c9a, points: 2, grade: true, solution: false -->
```python
assert(begaiement('abc', 3) == 'abbbccc')
```

<div class="alert alert-success">
    <h2>Les points à retenir</h2>
    
* accès à un caractère dans un chaîne par `[]`
* parcours d'une chaîne avec une boucle `for`
* immuabilité d'une chaîne de caractères
* `+` comme opérateur de concaténation de chaînes de caractères
* `*` comme opérateur de création d'une nouvelle chaîne
</div>